% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on eucliddiv_inv2
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for eucliddiv_inv2
% SZS output start Proof for eucliddiv_inv2
1. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($difference($difference(X1,r00),$product(q00,y00)),$product(X0,y00)) & ~0 = $sum($difference($difference($difference(X1,y00),r00),$product(q00,y00)),$product($sum(X0,1),y00))) [input]
2. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($sum($difference(X1,r00),$uminus($product(q00,y00))),$product(X0,y00)) & ~0 = $sum($sum($difference($difference(X1,y00),r00),$uminus($product(q00,y00))),$product($sum(X0,1),y00))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 143]
4. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 144]
6. $uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)) [theory axiom 147]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 143]
21. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X0,y00)) & ~0 = $sum($sum($sum($difference(X1,y00),$uminus(r00)),$uminus($product(q00,y00))),$product($sum(X0,1),y00))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int] : (0 = $sum($product($uminus(1),y00),X0) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X2,y00)) & ~0 = $sum($sum($sum($difference(X1,y00),$uminus(r00)),$uminus($product(q00,y00))),$product($sum(X2,1),y00))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int] : (0 = $sum($product($uminus(1),y00),X0) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X2,y00)) & 0 != $sum($sum($sum($difference(X1,y00),$uminus(r00)),$uminus($product(q00,y00))),$product($sum(X2,1),y00))) [flattening 22]
24. 0 != $sum($sum($sum($difference(sK1,y00),$uminus(r00)),$uminus($product(q00,y00))),$product($sum(sK2,1),y00)) [cnf transformation 23]
25. 0 = $sum($sum($sum(sK1,$uminus(r00)),$uminus($product(q00,y00))),$product(sK2,y00)) [cnf transformation 23]
26. 0 = $sum($product($uminus(1),y00),sK0) [cnf transformation 23]
28. 0 = $sum($product(-1,y00),sK0) [evaluation 26]
34. 1 <=> 0 = $sum($product(-1,y00),sK0) [avatar definition]
36. 0 = $sum($product(-1,y00),sK0) <- (1) [avatar component clause 34]
37. 1 [avatar split clause 28,34]
40. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 3,36]
43. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 3,36]
46. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 36,3]
49. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 36,3]
91. 2 <=> 0 = $sum(sK0,$product(-1,y00)) [avatar definition]
94. 2 | ~1 [avatar split clause 40,34,91]
95. 2 | ~1 [avatar split clause 43,34,91]
96. 2 | ~1 [avatar split clause 46,34,91]
97. 2 | ~1 [avatar split clause 49,34,91]
729. 0 = $sum($sum($sum(sK1,$uminus(r00)),$uminus($product(q00,y00))),$product(y00,sK2)) [forward demodulation 25,14]
730. 0 = $sum($sum(sK1,$sum($uminus(r00),$uminus($product(q00,y00)))),$product(y00,sK2)) [forward demodulation 729,4]
731. 0 = $sum($sum(sK1,$uminus($sum($product(q00,y00),r00))),$product(y00,sK2)) [forward demodulation 730,6]
732. 0 = $sum($sum(sK1,$uminus($sum(r00,$product(q00,y00)))),$product(y00,sK2)) [forward demodulation 731,3]
821. 3 <=> 0 = $sum($sum(sK1,$uminus($sum(r00,$product(q00,y00)))),$product(y00,sK2)) [avatar definition]
824. 3 [avatar split clause 732,821]
938. 0 != $sum($sum($sum($difference(sK1,y00),$uminus(r00)),$uminus($product(q00,y00))),$product(y00,$sum(sK2,1))) [forward demodulation 24,14]
939. 0 != $sum($sum($sum($difference(sK1,y00),$uminus(r00)),$uminus($product(q00,y00))),$product(y00,$sum(1,sK2))) [forward demodulation 938,3]
940. 0 != $sum($sum($uminus($product(q00,y00)),$sum($difference(sK1,y00),$uminus(r00))),$product(y00,$sum(1,sK2))) [forward demodulation 939,3]
941. 0 != $sum($sum($uminus($product(q00,y00)),$sum($uminus(r00),$difference(sK1,y00))),$product(y00,$sum(1,sK2))) [forward demodulation 940,3]
1001. 4 <=> 0 = $sum($sum($uminus($product(q00,y00)),$sum($uminus(r00),$difference(sK1,y00))),$product(y00,$sum(1,sK2))) [avatar definition]
1004. ~4 [avatar split clause 941,1001]
1005. $false [avatar smt refutation 37,94,95,96,97,824,1004]
% SZS output end Proof for eucliddiv_inv2
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Refutation

% Memory used [KB]: 7803
% Time elapsed: 0.052 s
% ------------------------------
% ------------------------------
% Success in time 0.057 s
