% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on simpleloop_ind_inv1_hint
% Time limit reached!
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 13432
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_6 on simpleloop_ind_inv1_hint
% Time limit reached!
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 18805
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity:gtg=exists_all_2 on simpleloop_ind_inv1_hint
% Time limit reached!
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 12537
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity:thsq=on:thsqd=16_2 on simpleloop_ind_inv1_hint
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for simpleloop_ind_inv1_hint
% SZS output start Proof for simpleloop_ind_inv1_hint
1. x00 = x(0) & y00 = y(0) [input]
2. ! [X0 : $int] : ($greatereq(X0,0) => ($difference(x(X0),1) = x($sum(X0,1)) & $sum(y(X0),1) = y($sum(X0,1)) & $sum($difference($difference(x(X0),y(0)),x(0)),y(X0)) = inv(X0))) [input]
3. ~! [X0 : $int] : ($greatereq(X0,0) => 0 = inv(X0)) [input]
4. (0 = inv(0) & ! [X1 : $int] : ($greatereq(X1,0) => (0 = inv(X1) => 0 = inv($sum(X1,1))))) => ! [X0 : $int] : ($greatereq(X0,0) => 0 = inv(X0)) [input]
5. (0 = inv(0) & ! [X1 : $int] : (~$less(X1,0) => (0 = inv(X1) => 0 = inv($sum(X1,1))))) => ! [X0 : $int] : (~$less(X0,0) => 0 = inv(X0)) [theory normalization 4]
6. ~! [X0 : $int] : (~$less(X0,0) => 0 = inv(X0)) [theory normalization 3]
7. ! [X0 : $int] : (~$less(X0,0) => (x($sum(X0,1)) = $sum(x(X0),$uminus(1)) & $sum(y(X0),1) = y($sum(X0,1)) & inv(X0) = $sum($sum($difference(x(X0),y(0)),$uminus(x(0))),y(X0)))) [theory normalization 2]
8. $sum(X0,X1) = $sum(X1,X0) [theory axiom 143]
9. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 144]
10. $sum(X0,0) = X0 [theory axiom 145]
11. $uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)) [theory axiom 147]
12. 0 = $sum(X0,$uminus(X0)) [theory axiom 148]
13. ~$less(X0,X0) [theory axiom 150]
14. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) [theory axiom 151]
15. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom 152]
16. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom 153]
17. $less(X1,$sum(X0,1)) | $less(X0,X1) [theory axiom 155]
18. $uminus($uminus(X0)) = X0 [theory axiom 156]
19. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom 169]
20. ! [X0 : $int] : (~$less(X0,0) => (x($sum(X0,1)) = $sum(x(X0),$uminus(1)) & $sum(y(X0),1) = y($sum(X0,1)) & inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x(0))),y(X0)))) [theory normalization 7]
21. (0 = inv(0) & ! [X0 : $int] : (~$less(X0,0) => (0 = inv(X0) => 0 = inv($sum(X0,1))))) => ! [X1 : $int] : (~$less(X1,0) => 0 = inv(X1)) [rectify 5]
22. ! [X1 : $int] : (0 = inv(X1) | $less(X1,0)) | (0 != inv(0) | ? [X0 : $int] : ((0 != inv($sum(X0,1)) & 0 = inv(X0)) & ~$less(X0,0))) [ennf transformation 21]
23. ! [X1 : $int] : (0 = inv(X1) | $less(X1,0)) | 0 != inv(0) | ? [X0 : $int] : (0 != inv($sum(X0,1)) & 0 = inv(X0) & ~$less(X0,0)) [flattening 22]
24. ? [X0 : $int] : (0 != inv(X0) & ~$less(X0,0)) [ennf transformation 6]
25. ! [X0 : $int] : ((x($sum(X0,1)) = $sum(x(X0),$uminus(1)) & $sum(y(X0),1) = y($sum(X0,1)) & inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x(0))),y(X0))) | $less(X0,0)) [ennf transformation 20]
26. y00 = y(0) [cnf transformation 1]
27. x00 = x(0) [cnf transformation 1]
28. 0 != inv($sum(sK0,1)) | 0 != inv(0) | $less(X1,0) | 0 = inv(X1) [cnf transformation 23]
29. 0 = inv(sK0) | 0 != inv(0) | $less(X1,0) | 0 = inv(X1) [cnf transformation 23]
30. ~$less(sK0,0) | 0 != inv(0) | $less(X1,0) | 0 = inv(X1) [cnf transformation 23]
31. ~$less(sK1,0) [cnf transformation 24]
32. 0 != inv(sK1) [cnf transformation 24]
33. $less(X0,0) | inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x(0))),y(X0)) [cnf transformation 25]
34. $less(X0,0) | $sum(y(X0),1) = y($sum(X0,1)) [cnf transformation 25]
35. $less(X0,0) | x($sum(X0,1)) = $sum(x(X0),$uminus(1)) [cnf transformation 25]
36. $less(X0,0) | x($sum(X0,1)) = $sum(x(X0),-1) [evaluation 35]
38. 1 <=> $less(sK1,0) [avatar definition]
40. ~$less(sK1,0) <- (~1) [avatar component clause 38]
41. ~1 [avatar split clause 31,38]
43. 2 <=> y00 = y(0) [avatar definition]
45. y00 = y(0) <- (2) [avatar component clause 43]
46. 2 [avatar split clause 26,43]
48. 3 <=> x00 = x(0) [avatar definition]
50. x00 = x(0) <- (3) [avatar component clause 48]
51. 3 [avatar split clause 27,48]
53. 4 <=> 0 = inv(sK1) [avatar definition]
55. 0 != inv(sK1) <- (~4) [avatar component clause 53]
56. ~4 [avatar split clause 32,53]
57. $sum(y(sK1),1) = y($sum(sK1,1)) <- (~1) [resolution 34,40]
58. $sum(y(0),1) = y($sum(0,1)) [resolution 34,13]
59. $sum(y(0),1) = y(1) [evaluation 58]
60. y(1) = $sum(y00,1) <- (2) [forward demodulation 59,45]
62. 5 <=> $sum(y(sK1),1) = y($sum(sK1,1)) [avatar definition]
64. $sum(y(sK1),1) = y($sum(sK1,1)) <- (5) [avatar component clause 62]
65. 5 | 1 [avatar split clause 57,38,62]
66. 0 = $sum($uminus(X0),X0) [superposition 12,18]
68. 6 <=> y(1) = $sum(y00,1) [avatar definition]
70. y(1) = $sum(y00,1) <- (6) [avatar component clause 68]
71. 6 | ~2 [avatar split clause 60,43,68]
72. x($sum(sK1,1)) = $sum(x(sK1),-1) <- (~1) [resolution 36,40]
73. x($sum(0,1)) = $sum(x(0),-1) [resolution 36,13]
74. $sum(x(0),-1) = x(1) [evaluation 73]
75. x(1) = $sum(x00,-1) <- (3) [forward demodulation 74,50]
77. 7 <=> x($sum(sK1,1)) = $sum(x(sK1),-1) [avatar definition]
79. x($sum(sK1,1)) = $sum(x(sK1),-1) <- (7) [avatar component clause 77]
80. 7 | 1 [avatar split clause 72,38,77]
82. 8 <=> x(1) = $sum(x00,-1) [avatar definition]
84. x(1) = $sum(x00,-1) <- (8) [avatar component clause 82]
85. 8 | ~3 [avatar split clause 75,48,82]
87. 9 <=> ! [X1 : $int] : ($less(X1,0) | 0 = inv(X1)) [avatar definition]
88. $less(X1,0) | 0 = inv(X1) <- (9) [avatar component clause 87]
90. 10 <=> 0 = inv(0) [avatar definition]
91. 0 = inv(0) <- (10) [avatar component clause 90]
92. 0 != inv(0) <- (~10) [avatar component clause 90]
94. 11 <=> $less(sK0,0) [avatar definition]
95. $less(sK0,0) <- (11) [avatar component clause 94]
96. ~$less(sK0,0) <- (~11) [avatar component clause 94]
97. 9 | ~10 | ~11 [avatar split clause 30,94,90,87]
101. inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x00)),y(X0)) | $less(X0,0) <- (3) [forward demodulation 33,50]
102. $less(X0,0) | inv(X0) = $sum($sum($sum(x(X0),$uminus(y00)),$uminus(x00)),y(X0)) <- (2, 3) [forward demodulation 101,45]
103. inv(sK1) = $sum($sum($sum(x(sK1),$uminus(y00)),$uminus(x00)),y(sK1)) <- (~1, 2, 3) [resolution 102,40]
104. inv(0) = $sum($sum($sum(x(0),$uminus(y00)),$uminus(x00)),y(0)) <- (2, 3) [resolution 102,13]
105. inv(0) = $sum($sum($sum(x(0),$uminus(y00)),$uminus(x00)),y00) <- (2, 3) [forward demodulation 104,45]
106. inv(0) = $sum($sum($sum(x00,$uminus(y00)),$uminus(x00)),y00) <- (2, 3) [forward demodulation 105,50]
107. inv(X0) = $sum(y(X0),$sum($sum(x(X0),$uminus(y00)),$uminus(x00))) | $less(X0,0) <- (2, 3) [backward demodulation 102,8]
108. $less(X0,0) | y($sum(X0,1)) = $sum(1,y(X0)) [backward demodulation 34,8]
109. y($sum(sK1,1)) = $sum(1,y(sK1)) <- (5) [backward demodulation 64,8]
110. $sum(x(sK1),-1) = x($sum(1,sK1)) <- (7) [backward demodulation 79,8]
111. $less(X0,0) | x($sum(X0,1)) = $sum(-1,x(X0)) [backward demodulation 36,8]
112. $sum(0,X0) = X0 [superposition 8,10]
113. x(1) = $sum(-1,x00) <- (8) [superposition 8,84]
117. x(1) = $sum(-1,x00) <- (8) [superposition 8,84]
121. x(1) = $sum(-1,x00) <- (8) [superposition 84,8]
125. x(1) = $sum(-1,x00) <- (8) [superposition 84,8]
128. $less(X0,0) | inv(X0) = $sum(y(X0),$sum($uminus(x00),$sum(x(X0),$uminus(y00)))) <- (2, 3) [forward demodulation 107,8]
129. $sum(1,y(sK1)) = y($sum(1,sK1)) <- (5) [forward demodulation 109,8]
130. x($sum(1,sK1)) = $sum(-1,x(sK1)) <- (7) [forward demodulation 110,8]
138. 12 <=> $sum(1,y(sK1)) = y($sum(1,sK1)) [avatar definition]
141. 12 | ~5 [avatar split clause 129,62,138]
143. 13 <=> x($sum(1,sK1)) = $sum(-1,x(sK1)) [avatar definition]
145. x($sum(1,sK1)) = $sum(-1,x(sK1)) <- (13) [avatar component clause 143]
146. 13 | ~7 [avatar split clause 130,77,143]
148. 14 <=> x(1) = $sum(-1,x00) [avatar definition]
150. x(1) = $sum(-1,x00) <- (14) [avatar component clause 148]
151. 14 | ~8 [avatar split clause 113,82,148]
153. y($sum(0,1)) = $sum(1,y(0)) [resolution 108,13]
154. y(1) = $sum(1,y(0)) [evaluation 153]
156. y(1) = $sum(1,y00) <- (2) [forward demodulation 154,45]
158. x($sum(0,1)) = $sum(-1,x(0)) [resolution 111,13]
159. x(1) = $sum(-1,x(0)) [evaluation 158]
161. x(1) = $sum(-1,x00) <- (3) [forward demodulation 159,50]
163. 15 <=> y(1) = $sum(1,y00) [avatar definition]
165. y(1) = $sum(1,y00) <- (15) [avatar component clause 163]
166. 15 | ~2 [avatar split clause 156,43,163]
167. 14 | ~3 [avatar split clause 161,48,148]
168. $less(X0,$sum(X0,1)) [resolution 17,13]
171. $less(X4,y(1)) | $less(y00,X4) <- (6) [superposition 17,70]
172. $less(X5,1) | $less(0,X5) [superposition 17,112]
173. $less(X6,0) | $less($uminus(1),X6) [superposition 17,66]
174. $less(X6,0) | $less(-1,X6) [evaluation 173]
175. inv(sK1) = $sum($sum($uminus(x00),$sum(x(sK1),$uminus(y00))),y(sK1)) <- (~1, 2, 3) [forward demodulation 103,8]
177. 16 <=> inv(sK1) = $sum($sum($uminus(x00),$sum(x(sK1),$uminus(y00))),y(sK1)) [avatar definition]
179. inv(sK1) = $sum($sum($uminus(x00),$sum(x(sK1),$uminus(y00))),y(sK1)) <- (16) [avatar component clause 177]
180. 16 | 1 | ~2 | ~3 [avatar split clause 175,48,43,38,177]
181. $less(X0,$sum(1,X0)) [superposition 168,8]
183. $less(y00,y(1)) <- (6) [superposition 168,70]
187. 17 <=> $less(y00,y(1)) [avatar definition]
189. $less(y00,y(1)) <- (17) [avatar component clause 187]
190. 17 | ~6 [avatar split clause 183,68,187]
191. inv(sK1) = $sum(y(sK1),$sum($uminus(x00),$sum(x(sK1),$uminus(y00)))) <- (~1, 2, 3) [resolution 128,40]
199. ~$less(0,$sum(sK0,1)) <- (11) [resolution 19,95]
205. ~$less(0,$sum(1,sK0)) <- (11) [forward demodulation 199,8]
206. ~$less(X0,sK0) | $less(X0,0) <- (11) [resolution 14,95]
210. ~$less(X7,X8) | $less(X7,$sum(X9,1)) | $less(X9,X8) [resolution 14,17]
211. $less(X10,$sum(X11,1)) | ~$less(X10,X11) [resolution 14,168]
213. 18 <=> inv(sK1) = $sum(y(sK1),$sum($uminus(x00),$sum(x(sK1),$uminus(y00)))) [avatar definition]
215. inv(sK1) = $sum(y(sK1),$sum($uminus(x00),$sum(x(sK1),$uminus(y00)))) <- (18) [avatar component clause 213]
216. 18 | 1 | ~2 | ~3 [avatar split clause 191,48,43,38,213]
217. $less(X0,y(1)) | ~$less(X0,y00) <- (17) [resolution 189,14]
221. 19 <=> $less(0,$sum(1,sK0)) [avatar definition]
222. $less(0,$sum(1,sK0)) <- (19) [avatar component clause 221]
224. ~19 | ~11 [avatar split clause 205,94,221]
232. 14 | ~8 [avatar split clause 117,82,148]
233. 14 | ~8 [avatar split clause 121,82,148]
234. 14 | ~8 [avatar split clause 125,82,148]
235. $less(0,sK1) | 0 = sK1 <- (~1) [resolution 15,40]
249. 20 <=> 0 = sK1 [avatar definition]
253. 21 <=> $less(0,sK1) [avatar definition]
255. $less(0,sK1) <- (21) [avatar component clause 253]
256. 20 | 21 | 1 [avatar split clause 235,38,253,249]
258. ~$less(sK1,$sum(0,1)) <- (21) [resolution 255,19]
259. ~$less(sK1,1) <- (21) [evaluation 258]
261. 22 <=> $less(sK1,1) [avatar definition]
263. ~$less(sK1,1) <- (~22) [avatar component clause 261]
264. ~22 | ~21 [avatar split clause 259,253,261]
265. $less(1,sK1) | 1 = sK1 <- (~22) [resolution 263,15]
280. $less(-1,sK1) <- (~1) [resolution 174,40]
290. 23 <=> $less(-1,sK1) [avatar definition]
292. $less(-1,sK1) <- (23) [avatar component clause 290]
293. 23 | 1 [avatar split clause 280,38,290]
294. $less(X0,sK1) | ~$less(X0,-1) <- (23) [resolution 292,14]
315. 24 <=> 1 = sK1 [avatar definition]
319. 25 <=> $less(1,sK1) [avatar definition]
321. $less(1,sK1) <- (25) [avatar component clause 319]
322. 24 | 25 | 22 [avatar split clause 265,261,319,315]
323. $less(X0,sK1) | ~$less(X0,1) <- (25) [resolution 321,14]
324. ~$less(sK1,$sum(1,1)) <- (25) [resolution 321,19]
325. ~$less(sK1,2) <- (25) [evaluation 324]
327. 26 <=> $less(sK1,2) [avatar definition]
329. ~$less(sK1,2) <- (~26) [avatar component clause 327]
330. ~26 | ~25 [avatar split clause 325,319,327]
331. $less(2,sK1) | sK1 = 2 <- (~26) [resolution 329,15]
333. ~$less(sK1,-1) <- (23) [resolution 294,13]
343. 27 <=> $less(sK1,-1) [avatar definition]
346. ~27 | ~23 [avatar split clause 333,290,343]
347. ~$less(y(1),y00) <- (17) [resolution 217,13]
384. $less($sum(sK0,X15),$sum(0,X15)) <- (11) [resolution 16,95]
397. $less($sum(0,X42),$sum(sK1,X42)) <- (21) [resolution 16,255]
398. $less($sum(1,X43),$sum(sK1,X43)) <- (25) [resolution 16,321]
399. $less($sum(-1,X44),$sum(sK1,X44)) <- (23) [resolution 16,292]
401. $less($sum(sK0,X15),X15) <- (11) [forward demodulation 384,112]
406. $less(X42,$sum(sK1,X42)) <- (21) [forward demodulation 397,112]
408. ~$less(X2,$sum(sK0,X3)) | $less(X2,X3) <- (11) [resolution 401,14]
409. ~$less(X4,$sum($sum(sK0,X4),1)) <- (11) [resolution 401,19]
410. $less($sum(sK0,sK0),0) <- (11) [resolution 401,206]
412. $less(0,$uminus(sK0)) <- (11) [superposition 401,12]
415. ~$less(X4,$sum(1,$sum(sK0,X4))) <- (11) [forward demodulation 409,8]
417. $sum(X2,$sum($uminus(X2),X3)) = $sum(0,X3) [superposition 9,12]
432. 0 = $sum(X5,$sum(X6,$uminus($sum(X5,X6)))) [superposition 9,12]
433. $sum(X7,$sum(X8,X9)) = $sum(X9,$sum(X7,X8)) [superposition 9,8]
443. $sum(X2,$sum($uminus(X2),X3)) = X3 [forward demodulation 417,112]
452. 28 <=> $less(0,$uminus(sK0)) [avatar definition]
454. $less(0,$uminus(sK0)) <- (28) [avatar component clause 452]
455. 28 | ~11 [avatar split clause 412,94,452]
456. $less(-1,sK0) <- (~11) [resolution 96,174]
457. inv(sK0) = $sum(y(sK0),$sum($uminus(x00),$sum(x(sK0),$uminus(y00)))) <- (2, 3, ~11) [resolution 96,128]
458. x($sum(sK0,1)) = $sum(-1,x(sK0)) <- (~11) [resolution 96,111]
459. y($sum(sK0,1)) = $sum(1,y(sK0)) <- (~11) [resolution 96,108]
460. $less(0,sK0) | 0 = sK0 <- (~11) [resolution 96,15]
462. $sum(-1,x(sK0)) = x($sum(1,sK0)) <- (~11) [forward demodulation 458,8]
463. $sum(1,y(sK0)) = y($sum(1,sK0)) <- (~11) [forward demodulation 459,8]
465. 29 <=> $sum(-1,x(sK0)) = x($sum(1,sK0)) [avatar definition]
467. $sum(-1,x(sK0)) = x($sum(1,sK0)) <- (29) [avatar component clause 465]
468. 29 | 11 [avatar split clause 462,94,465]
477. 28 | ~11 [avatar split clause 412,94,452]
479. $less(X1,$uminus(sK0)) | ~$less(X1,0) <- (28) [resolution 454,14]
480. ~$less($uminus(sK0),$sum(0,1)) <- (28) [resolution 454,19]
481. ~$less($uminus(sK0),1) <- (28) [evaluation 480]
484. 30 <=> $less($uminus(sK0),1) [avatar definition]
485. $less($uminus(sK0),1) <- (30) [avatar component clause 484]
487. ~30 | ~28 [avatar split clause 481,452,484]
491. ~19 | ~11 [avatar split clause 205,94,221]
497. 31 <=> $less(y(1),y00) [avatar definition]
500. ~31 | ~17 [avatar split clause 347,187,497]
514. $less($sum(sK0,sK0),0) <- (11) [resolution 206,401]
519. ~$less($sum(sK1,X4),$sum(X4,1)) <- (21) [resolution 406,19]
521. $less($uminus(sK1),0) <- (21) [superposition 406,12]
522. $less(X0,$sum(X0,sK1)) <- (21) [superposition 406,8]
526. 32 <=> $less($uminus(sK1),0) [avatar definition]
528. $less($uminus(sK1),0) <- (32) [avatar component clause 526]
529. 32 | ~21 [avatar split clause 521,253,526]
532. ~$less(0,$sum($uminus(sK1),1)) <- (32) [resolution 528,19]
534. ~$less(0,$sum(1,$uminus(sK1))) <- (32) [forward demodulation 532,8]
536. 33 <=> $less($sum(sK0,sK0),0) [avatar definition]
538. $less($sum(sK0,sK0),0) <- (33) [avatar component clause 536]
539. 33 | ~11 [avatar split clause 410,94,536]
540. $less(-1,sK0) <- (~11) [resolution 96,174]
543. y($sum(sK0,1)) = $sum(1,y(sK0)) <- (~11) [resolution 96,108]
547. $sum(1,y(sK0)) = y($sum(1,sK0)) <- (~11) [forward demodulation 543,8]
549. 34 <=> $less(-1,sK0) [avatar definition]
551. $less(-1,sK0) <- (34) [avatar component clause 549]
552. 34 | 11 [avatar split clause 456,94,549]
562. 28 | ~11 [avatar split clause 412,94,452]
568. ~30 | ~28 [avatar split clause 481,452,484]
572. ~19 | ~11 [avatar split clause 205,94,221]
586. 33 | ~11 [avatar split clause 410,94,536]
597. 33 | ~11 [avatar split clause 514,94,536]
598. $less($sum($sum(sK0,sK0),X0),$sum(0,X0)) <- (33) [resolution 538,16]
600. ~$less(0,$sum($sum(sK0,sK0),1)) <- (33) [resolution 538,19]
601. $less($sum($sum(sK0,sK0),X0),X0) <- (33) [forward demodulation 598,112]
602. $less($sum(sK0,$sum(sK0,X0)),X0) <- (33) [forward demodulation 601,9]
603. ~$less(0,$sum(1,$sum(sK0,sK0))) <- (33) [forward demodulation 600,8]
645. 35 <=> $less(0,$sum(1,$uminus(sK1))) [avatar definition]
648. ~35 | ~32 [avatar split clause 534,526,645]
666. 36 <=> sK1 = 2 [avatar definition]
668. sK1 = 2 <- (36) [avatar component clause 666]
670. 37 <=> $less(2,sK1) [avatar definition]
672. $less(2,sK1) <- (37) [avatar component clause 670]
673. 36 | 37 | 26 [avatar split clause 331,327,670,666]
686. ~$less(sK1,$sum(2,1)) <- (37) [resolution 672,19]
687. ~$less(sK1,3) <- (37) [evaluation 686]
688. ~$less($uminus(sK0),0) <- (28) [resolution 479,13]
693. 38 <=> $less($uminus(sK0),0) [avatar definition]
695. ~$less($uminus(sK0),0) <- (~38) [avatar component clause 693]
696. ~38 | ~28 [avatar split clause 688,452,693]
700. y($sum(sK0,1)) = $sum(1,y(sK0)) <- (~11) [resolution 96,108]
704. $sum(1,y(sK0)) = y($sum(1,sK0)) <- (~11) [forward demodulation 700,8]
705. 34 | 11 [avatar split clause 456,94,549]
707. 39 <=> $sum(1,y(sK0)) = y($sum(1,sK0)) [avatar definition]
709. $sum(1,y(sK0)) = y($sum(1,sK0)) <- (39) [avatar component clause 707]
710. 39 | 11 [avatar split clause 463,94,707]
724. ~19 | ~11 [avatar split clause 205,94,221]
729. 28 | ~11 [avatar split clause 412,94,452]
735. ~30 | ~28 [avatar split clause 481,452,484]
749. 40 <=> $less(sK1,3) [avatar definition]
752. ~40 | ~37 [avatar split clause 687,670,749]
754. 0 != inv(2) <- (~4, 36) [backward demodulation 55,668]
756. x($sum(1,2)) = $sum(-1,x(2)) <- (13, 36) [backward demodulation 145,668]
776. $sum(-1,x(2)) = x(3) <- (13, 36) [evaluation 756]
780. 41 <=> 0 = inv(2) [avatar definition]
783. ~41 | 4 | ~36 [avatar split clause 754,666,53,780]
784. ~38 | ~28 [avatar split clause 688,452,693]
787. 42 <=> $sum(-1,x(2)) = x(3) [avatar definition]
790. 42 | ~13 | ~36 [avatar split clause 776,666,143,787]
839. $less(-1,$uminus(sK0)) <- (~38) [resolution 695,174]
858. 33 | ~11 [avatar split clause 410,94,536]
879. ~40 | ~37 [avatar split clause 687,670,749]
880. 33 | ~11 [avatar split clause 514,94,536]
883. ~$less(0,$sum($sum(sK0,sK0),1)) <- (33) [resolution 538,19]
886. ~$less(0,$sum(1,$sum(sK0,sK0))) <- (33) [forward demodulation 883,8]
914. 43 <=> $less(-1,$uminus(sK0)) [avatar definition]
917. 43 | 38 [avatar split clause 839,693,914]
955. $less($sum($uminus(1),X14),X14) [superposition 181,443]
962. $less($sum(-1,X14),X14) [evaluation 955]
1011. $less($sum(-1,sK0),0) <- (11) [resolution 962,206]
1012. $less(x(1),x00) <- (14) [superposition 962,150]
1022. 44 <=> $less(x(1),x00) [avatar definition]
1025. 44 | ~14 [avatar split clause 1012,148,1022]
1065. 45 <=> $less($sum(-1,sK0),0) [avatar definition]
1067. $less($sum(-1,sK0),0) <- (45) [avatar component clause 1065]
1068. 45 | ~11 [avatar split clause 1011,94,1065]
1074. $less(0,sK0) | 0 = sK0 <- (~11) [resolution 96,15]
1077. 39 | 11 [avatar split clause 463,94,707]
1078. 34 | 11 [avatar split clause 456,94,549]
1079. 34 | 11 [avatar split clause 540,94,549]
1080. 39 | 11 [avatar split clause 547,94,707]
1081. 39 | 11 [avatar split clause 704,94,707]
1083. $less(X1,sK0) | ~$less(X1,-1) <- (34) [resolution 551,14]
1090. 46 <=> inv(sK0) = $sum(y(sK0),$sum($uminus(x00),$sum(x(sK0),$uminus(y00)))) [avatar definition]
1091. inv(sK0) != $sum(y(sK0),$sum($uminus(x00),$sum(x(sK0),$uminus(y00)))) <- (~46) [avatar component clause 1090]
1093. 46 | ~2 | ~3 | 11 [avatar split clause 457,94,48,43,1090]
1113. ~19 | ~11 [avatar split clause 205,94,221]
1118. 28 | ~11 [avatar split clause 412,94,452]
1124. ~30 | ~28 [avatar split clause 481,452,484]
1138. ~38 | ~28 [avatar split clause 688,452,693]
1148. 33 | ~11 [avatar split clause 410,94,536]
1168. 33 | ~11 [avatar split clause 514,94,536]
1171. ~$less(0,$sum($sum(sK0,sK0),1)) <- (33) [resolution 538,19]
1174. ~$less(0,$sum(1,$sum(sK0,sK0))) <- (33) [forward demodulation 1171,8]
1175. 43 | 38 [avatar split clause 839,693,914]
1271. 45 | ~11 [avatar split clause 1011,94,1065]
1289. ~$less(0,$sum($sum(-1,sK0),1)) <- (45) [resolution 1067,19]
1290. ~$less(0,sK0) <- (45) [evaluation 1289]
1304. ~$less($sum(sK1,y00),y(1)) <- (6, 21) [superposition 519,70]
1307. ~$less($sum(sK1,$uminus(1)),0) <- (21) [superposition 519,66]
1308. ~$less($sum(sK1,-1),0) <- (21) [evaluation 1307]
1312. ~$less($sum(y00,sK1),y(1)) <- (6, 21) [forward demodulation 1304,8]
1314. ~$less($sum(-1,sK1),0) <- (21) [forward demodulation 1308,8]
1316. 47 <=> $less($sum(y00,sK1),y(1)) [avatar definition]
1318. ~$less($sum(y00,sK1),y(1)) <- (~47) [avatar component clause 1316]
1319. ~47 | ~6 | ~21 [avatar split clause 1312,253,68,1316]
1320. ~$less($sum(y00,sK1),y00) <- (17, ~47) [resolution 1318,217]
1325. 48 <=> $less(0,sK0) [avatar definition]
1327. ~$less(0,sK0) <- (~48) [avatar component clause 1325]
1328. ~48 | ~45 [avatar split clause 1290,1065,1325]
1345. $less($sum(sK0,$sum(sK0,sK0)),0) <- (11, 33) [resolution 602,206]
1350. $less($sum(sK0,0),$uminus(sK0)) <- (33) [superposition 602,12]
1354. $less(sK0,$uminus(sK0)) <- (33) [forward demodulation 1350,10]
1359. 49 <=> $less(sK0,$uminus(sK0)) [avatar definition]
1362. 49 | ~33 [avatar split clause 1354,536,1359]
1368. 50 <=> $less(0,$sum(1,$sum(sK0,sK0))) [avatar definition]
1369. $less(0,$sum(1,$sum(sK0,sK0))) <- (50) [avatar component clause 1368]
1371. ~50 | ~33 [avatar split clause 603,536,1368]
1373. 51 <=> $less($sum(-1,sK1),0) [avatar definition]
1376. ~51 | ~21 [avatar split clause 1314,253,1373]
1377. ~50 | ~33 [avatar split clause 886,536,1368]
1400. ~50 | ~33 [avatar split clause 1174,536,1368]
1421. 52 <=> $less($sum(y00,sK1),y00) [avatar definition]
1424. ~52 | ~17 | 47 [avatar split clause 1320,1316,187,1421]
1438. $less(0,$sum(sK1,$uminus(1))) <- (25) [superposition 398,12]
1445. $less(0,$sum(sK1,-1)) <- (25) [evaluation 1438]
1452. $less(0,$sum(-1,sK1)) <- (25) [forward demodulation 1445,8]
1483. ~$less($uminus($sum(1,sK0)),0) <- (11) [superposition 415,432]
1539. 53 <=> $less($uminus($sum(1,sK0)),0) [avatar definition]
1542. ~53 | ~11 [avatar split clause 1483,94,1539]
1555. ~$less($sum(sK1,X4),$sum($sum(-1,X4),1)) <- (23) [resolution 399,19]
1562. $less(0,$sum(sK1,$uminus(-1))) <- (23) [superposition 399,12]
1570. $less(0,$sum(sK1,1)) <- (23) [evaluation 1562]
1572. ~$less($sum(sK1,X4),X4) <- (23) [evaluation 1555]
1577. $less(0,$sum(1,sK1)) <- (23) [forward demodulation 1570,8]
1579. 54 <=> $less(0,$sum(-1,sK1)) [avatar definition]
1582. 54 | ~25 [avatar split clause 1452,319,1579]
1637. ~$less(0,$uminus(sK1)) <- (23) [superposition 1572,12]
1648. 55 <=> $less(0,$uminus(sK1)) [avatar definition]
1651. ~55 | ~23 [avatar split clause 1637,290,1648]
1676. $less(sK0,$sum(X39,1)) | $less(X39,0) <- (11) [resolution 210,95]
1696. $less(0,$sum(X77,1)) | $less(X77,sK1) <- (21) [resolution 210,255]
1724. 56 <=> $less(0,$sum(1,sK1)) [avatar definition]
1727. 56 | ~23 [avatar split clause 1577,290,1724]
1729. 57 <=> $less($sum(sK0,$sum(sK0,sK0)),0) [avatar definition]
1732. 57 | ~11 | ~33 [avatar split clause 1345,536,94,1729]
1765. $less(X0,1) | $less(sK0,X0) <- (11) [resolution 408,17]
1766. $less(sK0,1) <- (11) [resolution 408,168]
1768. $less(sK0,sK1) <- (11, 21) [resolution 408,522]
1790. 58 <=> $less(sK0,1) [avatar definition]
1791. ~$less(sK0,1) <- (~58) [avatar component clause 1790]
1793. 58 | ~11 [avatar split clause 1766,94,1790]
1838. $less(sK0,sK1) <- (11, ~22) [resolution 1765,263]
1852. 59 <=> $less(sK0,sK1) [avatar definition]
1855. 59 | ~11 | ~21 [avatar split clause 1768,253,94,1852]
1856. 59 | ~11 | 22 [avatar split clause 1838,261,94,1852]
1890. $less(sK0,y(1)) | $less(y00,0) <- (6, 11) [superposition 1676,70]
1903. 60 <=> $less(y00,0) [avatar definition]
1904. ~$less(y00,0) <- (~60) [avatar component clause 1903]
1905. $less(y00,0) <- (60) [avatar component clause 1903]
1907. 61 <=> $less(sK0,y(1)) [avatar definition]
1908. ~$less(sK0,y(1)) <- (~61) [avatar component clause 1907]
1909. $less(sK0,y(1)) <- (61) [avatar component clause 1907]
1910. 60 | 61 | ~6 | ~11 [avatar split clause 1890,94,68,1907,1903]
1912. $less($sum(y00,X1),$sum(0,X1)) <- (60) [resolution 1905,16]
1914. ~$less(0,$sum(y00,1)) <- (60) [resolution 1905,19]
1915. $less($sum(y00,X1),X1) <- (60) [forward demodulation 1912,112]
1916. ~$less(0,$sum(1,y00)) <- (60) [forward demodulation 1914,8]
1917. ~$less(0,y(1)) <- (15, 60) [forward demodulation 1916,165]
1918. ~$less(sK0,y00) <- (17, ~61) [resolution 1908,217]
1919. $less(y00,sK0) <- (6, ~61) [resolution 1908,171]
1924. 62 <=> $less(0,y(1)) [avatar definition]
1926. ~$less(0,y(1)) <- (~62) [avatar component clause 1924]
1927. ~62 | ~15 | ~60 [avatar split clause 1917,1903,163,1924]
1928. $less(-1,y00) <- (~60) [resolution 1904,174]
1932. $less(0,y00) | y00 = 0 <- (~60) [resolution 1904,15]
1941. ~$less(y(1),$sum(sK0,1)) <- (61) [resolution 1909,19]
1942. ~$less(y(1),$sum(1,sK0)) <- (61) [forward demodulation 1941,8]
1944. 63 <=> $less(-1,y00) [avatar definition]
1947. 63 | 60 [avatar split clause 1928,1903,1944]
1955. 64 <=> y00 = 0 [avatar definition]
1957. y00 = 0 <- (64) [avatar component clause 1955]
1959. 65 <=> $less(0,y00) [avatar definition]
1962. 64 | 65 | 60 [avatar split clause 1932,1903,1959,1955]
2050. ~62 | ~15 | ~60 [avatar split clause 1917,1903,163,1924]
2056. inv(sK1) = $sum(y(sK1),$sum(x(sK1),$sum($uminus(y00),$uminus(x00)))) <- (18) [backward demodulation 215,433]
2057. inv(sK1) = $sum($sum(x(sK1),$sum($uminus(y00),$uminus(x00))),y(sK1)) <- (16) [backward demodulation 179,433]
2058. inv(sK0) != $sum(y(sK0),$sum(x(sK0),$sum($uminus(y00),$uminus(x00)))) <- (~46) [backward demodulation 1091,433]
2181. inv(sK1) = $sum(y(sK1),$sum(x(sK1),$sum($uminus(x00),$uminus(y00)))) <- (18) [forward demodulation 2056,8]
2182. inv(sK1) = $sum(y(sK1),$sum(x(sK1),$uminus($sum(y00,x00)))) <- (18) [forward demodulation 2181,11]
2183. inv(sK1) = $sum(y(sK1),$sum(x(sK1),$uminus($sum(x00,y00)))) <- (18) [forward demodulation 2182,8]
2184. inv(sK1) = $sum($sum(x(sK1),$sum($uminus(x00),$uminus(y00))),y(sK1)) <- (16) [forward demodulation 2057,8]
2185. inv(sK1) = $sum($sum(x(sK1),$uminus($sum(y00,x00))),y(sK1)) <- (16) [forward demodulation 2184,11]
2186. inv(sK1) = $sum($sum(x(sK1),$uminus($sum(x00,y00))),y(sK1)) <- (16) [forward demodulation 2185,8]
2187. inv(sK0) != $sum(y(sK0),$sum(x(sK0),$sum($uminus(x00),$uminus(y00)))) <- (~46) [forward demodulation 2058,8]
2188. inv(sK0) != $sum(y(sK0),$sum(x(sK0),$uminus($sum(y00,x00)))) <- (~46) [forward demodulation 2187,11]
2189. inv(sK0) != $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) <- (~46) [forward demodulation 2188,8]
2230. 66 <=> inv(sK1) = $sum(y(sK1),$sum(x(sK1),$uminus($sum(x00,y00)))) [avatar definition]
2233. 66 | ~18 [avatar split clause 2183,213,2230]
2235. 67 <=> inv(sK1) = $sum($sum(x(sK1),$uminus($sum(x00,y00))),y(sK1)) [avatar definition]
2238. 67 | ~16 [avatar split clause 2186,177,2235]
2247. $less(y(1),1) <- (6, 60) [superposition 1915,70]
2255. $less(0,$uminus(y00)) <- (60) [superposition 1915,12]
2261. 68 <=> $less(y(1),1) [avatar definition]
2264. 68 | ~6 | ~60 [avatar split clause 2247,1903,68,2261]
2266. 69 <=> inv(sK0) = $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) [avatar definition]
2267. inv(sK0) = $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) <- (69) [avatar component clause 2266]
2269. ~69 | 46 [avatar split clause 2189,1090,2266]
2289. 70 <=> $less(0,$uminus(y00)) [avatar definition]
2292. 70 | ~60 [avatar split clause 2255,1903,2289]
2335. 71 <=> $less(y(1),$sum(1,sK0)) [avatar definition]
2338. ~71 | ~61 [avatar split clause 1942,1907,2335]
2339. ~$less(sK0,y00) <- (17, ~61) [resolution 1908,217]
2340. $less(y00,sK0) <- (6, ~61) [resolution 1908,171]
2345. 72 <=> $less(sK0,y00) [avatar definition]
2348. ~72 | ~17 | 61 [avatar split clause 1918,1907,187,2345]
2354. ~71 | ~61 [avatar split clause 1942,1907,2335]
2403. $less(0,y(1)) | $less(y00,sK1) <- (6, 21) [superposition 1696,70]
2412. $less(y00,sK1) <- (6, 21, ~62) [subsumption resolution 2403,1926]
2414. 73 <=> $less(y00,sK1) [avatar definition]
2415. ~$less(y00,sK1) <- (~73) [avatar component clause 2414]
2416. $less(y00,sK1) <- (73) [avatar component clause 2414]
2417. 73 | ~6 | ~21 | 62 [avatar split clause 2412,1924,253,68,2414]
2421. ~$less(sK1,$sum(y00,1)) <- (73) [resolution 2416,19]
2422. ~$less(sK1,$sum(1,y00)) <- (73) [forward demodulation 2421,8]
2423. ~$less(sK1,y(1)) <- (15, 73) [forward demodulation 2422,165]
2426. ~$less($sum(X3,1),X3) [resolution 211,13]
2454. x($sum($sum(0,1),1)) = $sum(-1,x($sum(0,1))) [resolution 2426,111]
2455. y($sum($sum(0,1),1)) = $sum(1,y($sum(0,1))) [resolution 2426,108]
2459. $less(sK0,$sum(1,1)) <- (11) [resolution 2426,1765]
2472. $less(sK0,2) <- (11) [evaluation 2459]
2476. y(2) = $sum(1,y(1)) [evaluation 2455]
2477. x(2) = $sum(-1,x(1)) [evaluation 2454]
2490. 74 <=> $less(sK0,2) [avatar definition]
2493. 74 | ~11 [avatar split clause 2472,94,2490]
2500. 75 <=> y(2) = $sum(1,y(1)) [avatar definition]
2503. 75 [avatar split clause 2476,2500]
2505. 76 <=> x(2) = $sum(-1,x(1)) [avatar definition]
2508. 76 [avatar split clause 2477,2505]
2510. 77 <=> $less(sK1,y(1)) [avatar definition]
2512. ~$less(sK1,y(1)) <- (~77) [avatar component clause 2510]
2513. ~77 | ~15 | ~73 [avatar split clause 2423,2414,163,2510]
2539. ~$less(y00,1) <- (25, ~73) [resolution 2415,323]
2540. ~$less(y00,-1) <- (23, ~73) [resolution 2415,294]
2550. 63 | 60 [avatar split clause 1928,1903,1944]
2574. inv(0) = $sum($sum($uminus(x00),$sum(x00,$uminus(y00))),y00) <- (2, 3) [forward demodulation 106,8]
2575. inv(0) = $sum($sum(x00,$sum($uminus(y00),$uminus(x00))),y00) <- (2, 3) [forward demodulation 2574,433]
2576. inv(0) = $sum($sum(x00,$sum($uminus(x00),$uminus(y00))),y00) <- (2, 3) [forward demodulation 2575,8]
2577. inv(0) = $sum($uminus(y00),y00) <- (2, 3) [forward demodulation 2576,443]
2578. 0 = inv(0) <- (2, 3) [forward demodulation 2577,66]
2579. $false <- (2, 3, ~10) [subsumption resolution 2578,92]
2580. ~2 | ~3 | 10 [avatar contradiction clause 2579]
2581. 0 != inv($sum(1,sK0)) | 0 != inv(0) | $less(X1,0) | 0 = inv(X1) [forward demodulation 28,8]
2589. 0 = inv($sum(sK1,0)) <- (9, 23) [resolution 88,1572]
2593. 0 = inv(sK1) <- (~1, 9) [resolution 88,40]
2596. 0 = inv(sK1) <- (9, 23) [forward demodulation 2589,10]
2597. $false <- (~4, 9, 23) [subsumption resolution 2596,55]
2598. 4 | ~9 | ~23 [avatar contradiction clause 2597]
2599. $false <- (~1, ~4, 9) [subsumption resolution 2593,55]
2600. 1 | 4 | ~9 [avatar contradiction clause 2599]
2606. 0 = sK0 <- (~11, ~48) [subsumption resolution 1074,1327]
2659. 78 <=> 0 = sK0 [avatar definition]
2660. 0 != sK0 <- (~78) [avatar component clause 2659]
2662. 78 | 11 | 48 [avatar split clause 2606,1325,94,2659]
2680. 34 | 11 [avatar split clause 456,94,549]
2681. 0 = inv(sK0) | $less(X1,0) | 0 = inv(X1) <- (10) [subsumption resolution 29,91]
2682. 34 | 11 [avatar split clause 540,94,549]
2689. 79 <=> 0 = inv(sK0) [avatar definition]
2691. 0 = inv(sK0) <- (79) [avatar component clause 2689]
2692. 9 | 79 | ~10 [avatar split clause 2681,90,2689,87]
2693. 0 = $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) <- (69, 79) [forward demodulation 2267,2691]
2694. $less(0,sK0) <- (~58) [resolution 1791,172]
2698. 80 <=> 0 = $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) [avatar definition]
2700. 0 = $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) <- (80) [avatar component clause 2698]
2701. 80 | ~69 | ~79 [avatar split clause 2693,2689,2266,2698]
2732. 0 != inv($sum(1,sK0)) | $less(X1,0) | 0 = inv(X1) <- (10) [subsumption resolution 2581,91]
2739. 81 <=> 0 = inv($sum(1,sK0)) [avatar definition]
2741. 0 != inv($sum(1,sK0)) <- (~81) [avatar component clause 2739]
2742. 9 | ~81 | ~10 [avatar split clause 2732,90,2739,87]
2748. $less(X2,$sum(1,sK0)) | ~$less(X2,0) <- (19) [resolution 222,14]
2749. ~$less($sum(1,sK0),$sum(0,1)) <- (19) [resolution 222,19]
2750. ~$less($sum(1,sK0),1) <- (19) [evaluation 2749]
2755. $less(0,sK0) <- (~11, ~78) [subsumption resolution 460,2660]
2756. 48 | 11 | 78 [avatar split clause 2755,2659,94,1325]
2762. ~$less(1,$sum($uminus(sK0),1)) <- (30) [resolution 485,19]
2763. ~$less(1,$sum(1,$uminus(sK0))) <- (30) [forward demodulation 2762,8]
2764. 48 | 58 [avatar split clause 2694,1790,1325]
2794. 82 <=> $less($sum(1,sK0),1) [avatar definition]
2797. ~82 | ~19 [avatar split clause 2750,221,2794]
2802. 83 <=> $less(y00,1) [avatar definition]
2805. ~83 | ~25 | 73 [avatar split clause 2539,2414,319,2802]
2813. 0 = y(0) <- (2, 64) [backward demodulation 45,1957]
2845. 84 <=> 0 = y(0) [avatar definition]
2846. 0 != y(0) <- (~84) [avatar component clause 2845]
2848. 84 | ~2 | ~64 [avatar split clause 2813,1955,43,2845]
2851. y00 != 0 <- (2, ~84) [superposition 2846,45]
2852. ~64 | ~2 | 84 [avatar split clause 2851,2845,43,1955]
2894. 73 | ~6 | ~21 | 62 [avatar split clause 2412,1924,253,68,2414]
2900. inv(0) = $sum(y(0),$sum($uminus(x00),$sum(x(0),$uminus(y00)))) <- (2, 3) [resolution 128,13]
2912. inv(0) = $sum(y(0),$sum(x(0),$sum($uminus(y00),$uminus(x00)))) <- (2, 3) [forward demodulation 2900,433]
2913. inv(0) = $sum(y(0),$sum(x(0),$sum($uminus(x00),$uminus(y00)))) <- (2, 3) [forward demodulation 2912,8]
2914. inv(0) = $sum(y(0),$sum(x(0),$uminus($sum(y00,x00)))) <- (2, 3) [forward demodulation 2913,11]
2915. inv(0) = $sum(y(0),$sum(x(0),$uminus($sum(x00,y00)))) <- (2, 3) [forward demodulation 2914,8]
2916. inv(0) = $sum(y(0),$sum(x00,$uminus($sum(x00,y00)))) <- (2, 3) [forward demodulation 2915,50]
2917. inv(0) = $sum(y00,$sum(x00,$uminus($sum(x00,y00)))) <- (2, 3) [forward demodulation 2916,45]
2918. 0 = $sum(y00,$sum(x00,$uminus($sum(x00,y00)))) <- (2, 3, 10) [forward demodulation 2917,91]
2956. ~72 | ~17 | 61 [avatar split clause 1918,1907,187,2345]
2968. 85 <=> $less(y00,sK0) [avatar definition]
2971. 85 | ~6 | 61 [avatar split clause 1919,1907,68,2968]
2972. ~62 | ~15 | ~60 [avatar split clause 1917,1903,163,1924]
3035. ~72 | ~17 | 61 [avatar split clause 2339,1907,187,2345]
3037. 85 | ~6 | 61 [avatar split clause 2340,1907,68,2968]
3062. 86 <=> 0 = $sum(y00,$sum(x00,$uminus($sum(x00,y00)))) [avatar definition]
3065. 86 | ~2 | ~3 | ~10 [avatar split clause 2918,90,48,43,3062]
3066. ~$less(sK0,-1) <- (34) [resolution 1083,13]
3078. 87 <=> $less(sK0,-1) [avatar definition]
3081. ~87 | ~34 [avatar split clause 3066,549,3078]
3102. 88 <=> $less(1,$sum(1,$uminus(sK0))) [avatar definition]
3105. ~88 | ~30 [avatar split clause 2763,484,3102]
3106. ~77 | ~15 | ~73 [avatar split clause 2423,2414,163,2510]
3175. ~$less(sK1,y00) <- (17, ~77) [resolution 217,2512]
3181. ~$less($sum(1,$sum(sK0,sK0)),$sum(0,1)) <- (50) [resolution 1369,19]
3182. ~$less($sum(1,$sum(sK0,sK0)),1) <- (50) [evaluation 3181]
3187. 89 <=> $less(sK1,y00) [avatar definition]
3190. ~89 | ~17 | 77 [avatar split clause 3175,2510,187,3187]
3282. 63 | 60 [avatar split clause 1928,1903,1944]
3316. ~83 | ~25 | 73 [avatar split clause 2539,2414,319,2802]
3371. 90 <=> $less(y00,-1) [avatar definition]
3374. ~90 | ~23 | 73 [avatar split clause 2540,2414,290,3371]
3392. 73 | ~6 | ~21 | 62 [avatar split clause 2412,1924,253,68,2414]
3443. ~62 | ~15 | ~60 [avatar split clause 1917,1903,163,1924]
3564. ~89 | ~17 | 77 [avatar split clause 3175,2510,187,3187]
3664. 91 <=> $less($sum(1,$sum(sK0,sK0)),1) [avatar definition]
3667. ~91 | ~50 [avatar split clause 3182,1368,3664]
3678. ~$less($sum(1,sK0),0) <- (19) [resolution 2748,13]
3693. 92 <=> $less($sum(1,sK0),0) [avatar definition]
3695. ~$less($sum(1,sK0),0) <- (~92) [avatar component clause 3693]
3696. ~92 | ~19 [avatar split clause 3678,221,3693]
3697. ~77 | ~15 | ~73 [avatar split clause 2423,2414,163,2510]
3698. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum($uminus(x00),$sum(x($sum(1,sK0)),$uminus(y00)))) <- (2, 3, ~92) [resolution 3695,128]
3706. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum($uminus(x00),$sum($uminus(y00),x($sum(1,sK0))))) <- (2, 3, ~92) [forward demodulation 3698,8]
3707. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum($uminus(x00),$sum($uminus(y00),$sum(-1,x(sK0))))) <- (2, 3, 29, ~92) [forward demodulation 3706,467]
3708. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum($uminus(x00),$sum(x(sK0),$sum($uminus(y00),-1)))) <- (2, 3, 29, ~92) [forward demodulation 3707,433]
3709. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(x(sK0),$sum($sum($uminus(y00),-1),$uminus(x00)))) <- (2, 3, 29, ~92) [forward demodulation 3708,433]
3710. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(x(sK0),$sum($uminus(x00),$sum($uminus(y00),-1)))) <- (2, 3, 29, ~92) [forward demodulation 3709,8]
3711. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(x(sK0),$sum(-1,$sum($uminus(x00),$uminus(y00))))) <- (2, 3, 29, ~92) [forward demodulation 3710,433]
3712. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(-1,$sum($sum($uminus(x00),$uminus(y00)),x(sK0)))) <- (2, 3, 29, ~92) [forward demodulation 3711,433]
3713. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(-1,$sum(x(sK0),$sum($uminus(x00),$uminus(y00))))) <- (2, 3, 29, ~92) [forward demodulation 3712,8]
3714. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(-1,$sum(x(sK0),$uminus($sum(y00,x00))))) <- (2, 3, 29, ~92) [forward demodulation 3713,11]
3715. inv($sum(1,sK0)) = $sum(y($sum(1,sK0)),$sum(-1,$sum(x(sK0),$uminus($sum(x00,y00))))) <- (2, 3, 29, ~92) [forward demodulation 3714,8]
3716. inv($sum(1,sK0)) = $sum($sum(1,y(sK0)),$sum(-1,$sum(x(sK0),$uminus($sum(x00,y00))))) <- (2, 3, 29, 39, ~92) [forward demodulation 3715,709]
3717. $sum(y(sK0),$sum(x(sK0),$uminus($sum(x00,y00)))) = inv($sum(1,sK0)) <- (2, 3, 29, 39, ~92) [evaluation 3716]
3718. 0 = inv($sum(1,sK0)) <- (2, 3, 29, 39, 80, ~92) [forward demodulation 3717,2700]
3719. $false <- (2, 3, 29, 39, 80, ~81, ~92) [subsumption resolution 3718,2741]
3720. ~2 | ~3 | ~29 | ~39 | ~80 | 81 | 92 [avatar contradiction clause 3719]
3725. $false [avatar smt refutation 41,46,51,56,65,71,80,85,97,141,146,151,166,167,180,190,216,224,232,233,234,256,264,293,322,330,346,455,468,477,487,491,500,529,539,552,562,568,572,586,597,648,673,696,705,710,724,729,735,752,783,784,790,858,879,880,917,1025,1068,1077,1078,1079,1080,1081,1093,1113,1118,1124,1138,1148,1168,1175,1271,1319,1328,1362,1371,1376,1377,1400,1424,1542,1582,1651,1727,1732,1793,1855,1856,1910,1927,1947,1962,2050,2233,2238,2264,2269,2292,2338,2348,2354,2417,2493,2503,2508,2513,2550,2580,2598,2600,2662,2680,2682,2692,2701,2742,2756,2764,2797,2805,2848,2852,2894,2956,2971,2972,3035,3037,3065,3081,3105,3106,3190,3282,3316,3374,3392,3443,3564,3667,3696,3697,3720]
% SZS output end Proof for simpleloop_ind_inv1_hint
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Refutation

% Memory used [KB]: 9338
% Time elapsed: 0.192 s
% ------------------------------
% ------------------------------
% Success in time 2.117 s
