% Warning: check-sat is not the last entry. Skipping the rest!
% lrs+10_1_iik=one:ind=int:sos=theory:sstl=1_87 on simpleloop_ind_inv1_hint
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for simpleloop_ind_inv1_hint
% SZS output start Proof for simpleloop_ind_inv1_hint
1. x00 = x(0) & y00 = y(0) [input]
2. ! [X0 : $int] : ($greatereq(X0,0) => ($difference(x(X0),1) = x($sum(X0,1)) & $sum(y(X0),1) = y($sum(X0,1)) & $sum($difference($difference(x(X0),y(0)),x(0)),y(X0)) = inv(X0))) [input]
3. ~! [X0 : $int] : ($greatereq(X0,0) => 0 = inv(X0)) [input]
6. ~! [X0 : $int] : (~$less(X0,0) => 0 = inv(X0)) [theory normalization 3]
7. ! [X0 : $int] : (~$less(X0,0) => (x($sum(X0,1)) = $sum(x(X0),$uminus(1)) & $sum(y(X0),1) = y($sum(X0,1)) & inv(X0) = $sum($sum($difference(x(X0),y(0)),$uminus(x(0))),y(X0)))) [theory normalization 2]
8. $sum(X0,X1) = $sum(X1,X0) [theory axiom 143]
9. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 144]
10. $sum(X0,0) = X0 [theory axiom 145]
11. $uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)) [theory axiom 147]
12. 0 = $sum(X0,$uminus(X0)) [theory axiom 148]
13. ~$less(X0,X0) [theory axiom 150]
14. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) [theory axiom 151]
15. $less(X1,X0) | $less(X0,X1) | X0 = X1 [theory axiom 152]
16. $less($sum(X0,X2),$sum(X1,X2)) | ~$less(X0,X1) [theory axiom 153]
17. $less(X1,$sum(X0,1)) | $less(X0,X1) [theory axiom 155]
18. $uminus($uminus(X0)) = X0 [theory axiom 156]
20. ! [X0 : $int] : (~$less(X0,0) => (x($sum(X0,1)) = $sum(x(X0),$uminus(1)) & $sum(y(X0),1) = y($sum(X0,1)) & inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x(0))),y(X0)))) [theory normalization 7]
24. ? [X0 : $int] : (0 != inv(X0) & ~$less(X0,0)) [ennf transformation 6]
25. ! [X0 : $int] : ((x($sum(X0,1)) = $sum(x(X0),$uminus(1)) & $sum(y(X0),1) = y($sum(X0,1)) & inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x(0))),y(X0))) | $less(X0,0)) [ennf transformation 20]
29. ? [X0 : $int] : (0 != inv(X0) & ~$less(X0,0)) => (0 != inv(sK1) & ~$less(sK1,0)) [choice axiom]
30. 0 != inv(sK1) & ~$less(sK1,0) [skolemisation 24,29]
31. y00 = y(0) [cnf transformation 1]
32. x00 = x(0) [cnf transformation 1]
36. ~$less(sK1,0) [cnf transformation 30]
37. 0 != inv(sK1) [cnf transformation 30]
38. inv(X0) = $sum($sum($sum(x(X0),$uminus(y(0))),$uminus(x(0))),y(X0)) | $less(X0,0) [cnf transformation 25]
39. $sum(y(X0),1) = y($sum(X0,1)) | $less(X0,0) [cnf transformation 25]
40. x($sum(X0,1)) = $sum(x(X0),$uminus(1)) | $less(X0,0) [cnf transformation 25]
41. x($sum(X0,1)) = $sum(x(X0),-1) | $less(X0,0) [evaluation 40]
44. 1 <=> 0 = inv(0) [avatar definition]
45. 0 = inv(0) <- (1) [avatar component clause 44]
65. $less(X0,0) | x($sum(X0,1)) = $sum(-1,x(X0)) [forward demodulation 41,8]
66. $less(X0,0) | y($sum(X0,1)) = $sum(1,y(X0)) [forward demodulation 39,8]
67. inv(X0) = $sum(y(X0),$sum($sum(x(X0),$uminus(y(0))),$uminus(x(0)))) | $less(X0,0) [forward demodulation 38,8]
68. inv(X0) = $sum(y(X0),$sum($uminus(x(0)),$sum(x(X0),$uminus(y(0))))) | $less(X0,0) [forward demodulation 67,8]
69. inv(X0) = $sum(y(X0),$sum($uminus(x(0)),$sum(x(X0),$uminus(y00)))) | $less(X0,0) [forward demodulation 68,31]
70. $less(X0,0) | inv(X0) = $sum(y(X0),$sum($uminus(x00),$sum(x(X0),$uminus(y00)))) [forward demodulation 69,32]
71. (! [X0 : $int] : ((0 = inv(X0) & ~$less(X0,0)) => 0 = inv($sum(X0,1))) & 0 = inv(0)) => ! [X1 : $int] : (~$less(X1,0) => 0 = inv(X1)) [integer induction hypothesis (up, infinite interval)]
72. ! [X1 : $int] : (0 = inv(X1) | $less(X1,0)) | (? [X0 : $int] : (0 != inv($sum(X0,1)) & (0 = inv(X0) & ~$less(X0,0))) | 0 != inv(0)) [ennf transformation 71]
73. 0 != inv(0) | ~$less(sK3,0) | $less(X1,0) | 0 = inv(X1) [cnf transformation 72]
74. 0 != inv(0) | 0 = inv(sK3) | $less(X1,0) | 0 = inv(X1) [cnf transformation 72]
75. 0 != inv(0) | 0 != inv($sum(sK3,1)) | $less(X1,0) | 0 = inv(X1) [cnf transformation 72]
76. 0 != inv(0) | 0 != inv($sum(sK3,1)) | $less(sK1,0) [resolution 75,37]
77. 0 != inv(0) | 0 = inv(sK3) | $less(sK1,0) [resolution 74,37]
78. 0 != inv(0) | ~$less(sK3,0) | $less(sK1,0) [resolution 73,37]
79. 0 = $sum($uminus(X0),X0) [superposition 12,18]
90. $less(X0,$sum(X0,1)) [resolution 17,13]
170. $less(0,$sum(X3,$uminus(X2))) | ~$less(X2,X3) [superposition 16,12]
201. $sum($uminus(X18),$sum($uminus(X19),X20)) = $sum($uminus($sum(X19,X18)),X20) [superposition 9,11]
203. 0 = $sum(X5,$sum(X6,$uminus($sum(X5,X6)))) [superposition 9,12]
204. $sum(X7,$sum(X8,X9)) = $sum(X9,$sum(X7,X8)) [superposition 9,8]
252. inv(0) = $sum(y(0),$sum($uminus(x00),$sum(x(0),$uminus(y00)))) [resolution 70,13]
253. ~$less(X1,X0) | inv(X0) = $sum(y(X0),$sum($uminus(x00),$sum(x(X0),$uminus(y00)))) | $less(X1,0) [resolution 70,14]
254. inv(0) = $sum(y(0),$sum(x(0),$sum($uminus(y00),$uminus(x00)))) [forward demodulation 252,204]
255. inv(0) = $sum(y(0),$sum(x(0),$sum($uminus(x00),$uminus(y00)))) [forward demodulation 254,8]
256. inv(0) = $sum(y(0),$sum(x(0),$uminus($sum(y00,x00)))) [forward demodulation 255,11]
257. inv(0) = $sum(y(0),$sum(x(0),$uminus($sum(x00,y00)))) [forward demodulation 256,8]
258. inv(0) = $sum(y(0),$sum(x00,$uminus($sum(x00,y00)))) [forward demodulation 257,32]
259. inv(0) = $sum(y00,$sum(x00,$uminus($sum(x00,y00)))) [forward demodulation 258,31]
633. inv(0) = $sum(x00,$sum($uminus($sum(x00,y00)),y00)) [superposition 204,259]
680. inv(0) = $sum(x00,$sum(y00,$uminus($sum(x00,y00)))) [forward demodulation 633,8]
681. 0 = inv(0) [forward demodulation 680,203]
689. 1 [avatar split clause 681,44]
691. 0 != inv($sum(1,sK3)) | 0 != inv(0) | $less(sK1,0) [forward demodulation 76,8]
695. 14 <=> 0 = inv($sum(1,sK3)) [avatar definition]
697. 0 != inv($sum(1,sK3)) <- (~14) [avatar component clause 695]
702. 15 <=> 0 = inv(sK3) [avatar definition]
704. 0 = inv(sK3) <- (15) [avatar component clause 702]
709. 16 <=> $less(sK3,0) [avatar definition]
711. ~$less(sK3,0) <- (~16) [avatar component clause 709]
731. ~$less(sK3,0) | $less(sK1,0) <- (1) [subsumption resolution 78,45]
733. 17 <=> $less(sK1,0) [avatar definition]
736. 17 | ~16 | ~1 [avatar split clause 731,44,709,733]
737. 0 = inv(sK3) | $less(sK1,0) <- (1) [subsumption resolution 77,45]
738. 17 | 15 | ~1 [avatar split clause 737,44,702,733]
739. 0 != inv($sum(1,sK3)) | $less(sK1,0) <- (1) [subsumption resolution 691,45]
740. 17 | ~14 | ~1 [avatar split clause 739,44,695,733]
741. ~17 [avatar split clause 36,733]
770. y($sum(sK3,1)) = $sum(1,y(sK3)) <- (~16) [resolution 711,66]
771. x($sum(sK3,1)) = $sum(-1,x(sK3)) <- (~16) [resolution 711,65]
779. $sum(1,y(sK3)) = y($sum(1,sK3)) <- (~16) [forward demodulation 770,8]
780. $sum(-1,x(sK3)) = x($sum(1,sK3)) <- (~16) [forward demodulation 771,8]
782. 20 <=> 0 = sK3 [avatar definition]
784. 0 = sK3 <- (20) [avatar component clause 782]
786. 21 <=> $less(0,sK3) [avatar definition]
788. $less(0,sK3) <- (21) [avatar component clause 786]
2420. inv($sum(X9,1)) = $sum(y($sum(X9,1)),$sum($uminus(x00),$sum(x($sum(X9,1)),$uminus(y00)))) | $less(X9,0) [resolution 253,90]
2425. inv($sum(X15,$uminus(X16))) = $sum(y($sum(X15,$uminus(X16))),$sum($uminus(x00),$sum(x($sum(X15,$uminus(X16))),$uminus(y00)))) | $less(0,0) | ~$less(X16,X15) [resolution 253,170]
2445. inv($sum(X15,$uminus(X16))) = $sum(y($sum(X15,$uminus(X16))),$sum($uminus(x00),$sum(x($sum(X15,$uminus(X16))),$uminus(y00)))) | ~$less(X16,X15) [evaluation 2425]
2452. inv($sum(X9,1)) = $sum($uminus(x00),$sum($sum(x($sum(X9,1)),$uminus(y00)),y($sum(X9,1)))) | $less(X9,0) [forward demodulation 2420,204]
2453. inv($sum(X9,1)) = $sum($uminus(x00),$sum(y($sum(X9,1)),$sum(x($sum(X9,1)),$uminus(y00)))) | $less(X9,0) [forward demodulation 2452,8]
2454. inv($sum(X9,1)) = $sum($uminus(x00),$sum($uminus(y00),$sum(y($sum(X9,1)),x($sum(X9,1))))) | $less(X9,0) [forward demodulation 2453,204]
2455. inv($sum(X9,1)) = $sum($uminus($sum(y00,x00)),$sum(y($sum(X9,1)),x($sum(X9,1)))) | $less(X9,0) [forward demodulation 2454,201]
2456. inv($sum(X9,1)) = $sum($uminus($sum(y00,x00)),$sum(x($sum(X9,1)),y($sum(X9,1)))) | $less(X9,0) [forward demodulation 2455,8]
2457. $less(X9,0) | inv($sum(X9,1)) = $sum($uminus($sum(x00,y00)),$sum(x($sum(X9,1)),y($sum(X9,1)))) [forward demodulation 2456,8]
2470. inv($sum(X15,$uminus(X16))) = $sum($uminus(x00),$sum($sum(x($sum(X15,$uminus(X16))),$uminus(y00)),y($sum(X15,$uminus(X16))))) | ~$less(X16,X15) [forward demodulation 2445,204]
2471. inv($sum(X15,$uminus(X16))) = $sum($uminus(x00),$sum(y($sum(X15,$uminus(X16))),$sum(x($sum(X15,$uminus(X16))),$uminus(y00)))) | ~$less(X16,X15) [forward demodulation 2470,8]
2472. inv($sum(X15,$uminus(X16))) = $sum($uminus(x00),$sum($uminus(y00),$sum(y($sum(X15,$uminus(X16))),x($sum(X15,$uminus(X16)))))) | ~$less(X16,X15) [forward demodulation 2471,204]
2473. inv($sum(X15,$uminus(X16))) = $sum($uminus($sum(y00,x00)),$sum(y($sum(X15,$uminus(X16))),x($sum(X15,$uminus(X16))))) | ~$less(X16,X15) [forward demodulation 2472,201]
2474. inv($sum(X15,$uminus(X16))) = $sum($uminus($sum(y00,x00)),$sum(x($sum(X15,$uminus(X16))),y($sum(X15,$uminus(X16))))) | ~$less(X16,X15) [forward demodulation 2473,8]
2475. ~$less(X16,X15) | inv($sum(X15,$uminus(X16))) = $sum($uminus($sum(x00,y00)),$sum(x($sum(X15,$uminus(X16))),y($sum(X15,$uminus(X16))))) [forward demodulation 2474,8]
26363. $less(0,sK3) | 0 = sK3 <- (~16) [resolution 711,15]
43893. 1380 <=> 0 = inv(1) [avatar definition]
102013. inv($sum(sK3,1)) = $sum($uminus($sum(x00,y00)),$sum(x($sum(sK3,1)),y($sum(sK3,1)))) <- (~16) [resolution 2457,711]
102093. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum(x($sum(1,sK3)),y($sum(1,sK3)))) <- (~16) [forward demodulation 102013,8]
102094. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum(x($sum(1,sK3)),$sum(1,y(sK3)))) <- (~16) [forward demodulation 102093,779]
102095. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum($sum(1,y(sK3)),x($sum(1,sK3)))) <- (~16) [forward demodulation 102094,8]
102096. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum(1,$sum(y(sK3),x($sum(1,sK3))))) <- (~16) [forward demodulation 102095,9]
102097. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum(1,$sum(y(sK3),$sum(-1,x(sK3))))) <- (~16) [forward demodulation 102096,780]
102098. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum(y(sK3),x(sK3))) <- (~16) [evaluation 102097]
102099. inv($sum(1,sK3)) = $sum($uminus($sum(x00,y00)),$sum(x(sK3),y(sK3))) <- (~16) [forward demodulation 102098,8]
107283. inv($sum(sK3,$uminus(0))) = $sum($uminus($sum(x00,y00)),$sum(x($sum(sK3,$uminus(0))),y($sum(sK3,$uminus(0))))) <- (21) [resolution 2475,788]
107402. inv($sum(sK3,0)) = $sum($uminus($sum(x00,y00)),$sum(x($sum(sK3,0)),y($sum(sK3,0)))) <- (21) [evaluation 107283]
107490. inv(sK3) = $sum($uminus($sum(x00,y00)),$sum(x(sK3),y(sK3))) <- (21) [forward demodulation 107402,10]
107491. 0 = $sum($uminus($sum(x00,y00)),$sum(x(sK3),y(sK3))) <- (15, 21) [forward demodulation 107490,704]
107492. 0 = inv($sum(1,sK3)) <- (15, ~16, 21) [backward demodulation 102099,107491]
107493. $false <- (~14, 15, ~16, 21) [subsumption resolution 107492,697]
107494. 14 | ~15 | 16 | ~21 [avatar contradiction clause 107493]
107660. 20 | 21 | 16 [avatar split clause 26363,709,786,782]
107884. 0 != inv($sum(1,0)) <- (~14, 20) [backward demodulation 697,784]
107918. inv($sum(1,0)) = $sum($uminus($sum(x00,y00)),$sum(x(0),y(0))) <- (~16, 20) [backward demodulation 102099,784]
107924. inv(1) = $sum($uminus($sum(x00,y00)),$sum(x(0),y(0))) <- (~16, 20) [evaluation 107918]
107929. 0 != inv(1) <- (~14, 20) [evaluation 107884]
107993. inv(1) = $sum($uminus($sum(x00,y00)),$sum(x(0),y00)) <- (~16, 20) [forward demodulation 107924,31]
107994. inv(1) = $sum($uminus($sum(x00,y00)),$sum(y00,x(0))) <- (~16, 20) [forward demodulation 107993,8]
107995. $sum($uminus($sum(x00,y00)),$sum(y00,x00)) = inv(1) <- (~16, 20) [forward demodulation 107994,32]
108132. inv(1) = $sum($uminus($sum(x00,y00)),$sum(x00,y00)) <- (~16, 20) [forward demodulation 107995,8]
108133. 0 = inv(1) <- (~16, 20) [forward demodulation 108132,79]
108134. 1380 | 16 | ~20 [avatar split clause 108133,782,709,43893]
108135. ~1380 | 14 | ~20 [avatar split clause 107929,782,695,43893]
110089. $false [avatar sat refutation 689,736,738,740,741,107494,107660,108134,108135]
% SZS output end Proof for simpleloop_ind_inv1_hint
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Refutation

% Memory used [KB]: 61790
% Time elapsed: 3.668 s
% ------------------------------
% ------------------------------
% Success in time 3.667 s
