% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on eucliddiv_denorm_inv1_len3_4
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for eucliddiv_denorm_inv1_len3_4
% SZS output start Proof for eucliddiv_denorm_inv1_len3_4
1. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($difference($difference(X1,r00),$product(q00,y00)),$product(X0,y00)) & ~0 = $product($sum($product($uminus(1),y00),y00),$sum($sum($product($product($product($product($product($uminus(1),$difference(X1,y00)),$difference(X1,y00)),y00),y00),y00),$product($sum(X0,1),y00)),$product($product($product($difference(X1,y00),$difference(X1,y00)),y00),y00)))) [input]
2. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($sum($difference(X1,r00),$uminus($product(q00,y00))),$product(X0,y00)) & ~0 = $product($sum($product($uminus(1),y00),y00),$sum($sum($product($product($product($product($product($uminus(1),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00),y00),$product($sum(X0,1),y00)),$product($product($product($sum(X1,$uminus(y00)),$sum(X1,$uminus(y00))),y00),y00)))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 141]
4. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 142]
6. $uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)) [theory axiom 145]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 141]
18. $product(X0,$sum(X1,X2)) = $sum($product(X0,X1),$product(X0,X2)) [theory axiom 156]
21. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X0,y00)) & ~0 = $product($sum($product($uminus(1),y00),y00),$sum($sum($product($product($product($product($product($uminus(1),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00),y00),$product($sum(X0,1),y00)),$product($product($product($sum(X1,$uminus(y00)),$sum(X1,$uminus(y00))),y00),y00)))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int] : (0 = $sum($product($uminus(1),y00),X0) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X2,y00)) & ~0 = $product($sum($product($uminus(1),y00),y00),$sum($sum($product($product($product($product($product($uminus(1),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00),y00),$product($sum(X2,1),y00)),$product($product($product($sum(X1,$uminus(y00)),$sum(X1,$uminus(y00))),y00),y00)))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int] : (0 = $sum($product($uminus(1),y00),X0) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X2,y00)) & 0 != $product($sum($product($uminus(1),y00),y00),$sum($sum($product($product($product($product($product($uminus(1),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00),y00),$product($sum(X2,1),y00)),$product($product($product($sum(X1,$uminus(y00)),$sum(X1,$uminus(y00))),y00),y00)))) [flattening 22]
24. 0 != $product($sum($product($uminus(1),y00),y00),$sum($sum($product($product($product($product($product($uminus(1),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),y00),$product($sum(sK2,1),y00)),$product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00))) [cnf transformation 23]
25. 0 = $sum($sum($sum(sK1,$uminus(r00)),$uminus($product(q00,y00))),$product(sK2,y00)) [cnf transformation 23]
26. 0 = $sum($product($uminus(1),y00),sK0) [cnf transformation 23]
28. 0 != $product($sum($product(-1,y00),y00),$sum($sum($product($product($product($product($product(-1,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),y00),$product($sum(sK2,1),y00)),$product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00))) [evaluation 24]
29. 0 = $sum($product(-1,y00),sK0) [evaluation 26]
35. 1 <=> 0 = $sum($product(-1,y00),sK0) [avatar definition]
37. 0 = $sum($product(-1,y00),sK0) <- (1) [avatar component clause 35]
38. 1 [avatar split clause 29,35]
41. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 3,37]
44. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 3,37]
47. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 37,3]
50. 0 = $sum(sK0,$product(-1,y00)) <- (1) [superposition 37,3]
92. 2 <=> 0 = $sum(sK0,$product(-1,y00)) [avatar definition]
95. 2 | ~1 [avatar split clause 41,35,92]
96. 2 | ~1 [avatar split clause 44,35,92]
97. 2 | ~1 [avatar split clause 47,35,92]
98. 2 | ~1 [avatar split clause 50,35,92]
730. 0 = $sum($sum($sum(sK1,$uminus(r00)),$uminus($product(q00,y00))),$product(y00,sK2)) [forward demodulation 25,14]
731. 0 = $sum($sum(sK1,$sum($uminus(r00),$uminus($product(q00,y00)))),$product(y00,sK2)) [forward demodulation 730,4]
732. 0 = $sum($sum(sK1,$uminus($sum($product(q00,y00),r00))),$product(y00,sK2)) [forward demodulation 731,6]
733. 0 = $sum($sum(sK1,$uminus($sum(r00,$product(q00,y00)))),$product(y00,sK2)) [forward demodulation 732,3]
822. 3 <=> 0 = $sum($sum(sK1,$uminus($sum(r00,$product(q00,y00)))),$product(y00,sK2)) [avatar definition]
825. 3 [avatar split clause 733,822]
939. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product($product($product($product(-1,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),y00),$product($sum(sK2,1),y00)))) [forward demodulation 28,3]
940. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($sum(sK2,1),y00),$product($product($product($product($product(-1,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),y00)))) [forward demodulation 939,3]
941. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($sum(sK2,1),y00),$product(y00,$product($product($product($product(-1,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00))))) [forward demodulation 940,14]
942. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($sum(sK2,1),y00),$product(y00,$product(y00,$product($product($product(-1,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00)))))) [forward demodulation 941,14]
943. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($sum(sK2,1),y00),$product(y00,$product(y00,$product(y00,$product($product(-1,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))))))))) [forward demodulation 942,14]
944. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($sum(sK2,1),y00),$product(y00,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00)))))))))) [forward demodulation 943,14]
945. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$sum($product(y00,$sum(sK2,1)),$product(y00,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00)))))))))) [forward demodulation 944,14]
946. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$product(y00,$sum($sum(sK2,1),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00)))))))))) [forward demodulation 945,18]
947. 0 != $product($sum($product(-1,y00),y00),$sum($product($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),y00),$product(y00,$sum(sK2,$sum(1,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00))))))))))) [forward demodulation 946,4]
948. 0 != $product($sum($product(-1,y00),y00),$sum($product(y00,$product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00)),$product(y00,$sum(sK2,$sum(1,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00))))))))))) [forward demodulation 947,14]
949. 0 != $product($sum($product(-1,y00),y00),$product(y00,$sum($product($product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00))),y00),$sum(sK2,$sum(1,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00))))))))))) [forward demodulation 948,18]
950. 0 != $product($sum($product(-1,y00),y00),$product(y00,$sum($product(y00,$product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00)))),$sum(sK2,$sum(1,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00))))))))))) [forward demodulation 949,14]
951. 0 != $product($sum(y00,$product(-1,y00)),$product(y00,$sum($product(y00,$product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00)))),$sum(sK2,$sum(1,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00))))))))))) [forward demodulation 950,3]
1011. 4 <=> 0 = $product($sum(y00,$product(-1,y00)),$product(y00,$sum($product(y00,$product($sum(sK1,$uminus(y00)),$sum(sK1,$uminus(y00)))),$sum(sK2,$sum(1,$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(-1,$sum(sK1,$uminus(y00))))))))))) [avatar definition]
1014. ~4 [avatar split clause 951,1011]
1015. $false [avatar sat refutation 38,95,96,97,98,825,1014]
% SZS output end Proof for eucliddiv_denorm_inv1_len3_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Refutation

% Memory used [KB]: 7547
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
% Success in time 0.04 s
