% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on cubes_denorm_inv4_len5_5
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_6 on cubes_denorm_inv4_len5_5
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8187
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity:gtg=exists_all_2 on cubes_denorm_inv4_len5_5
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for cubes_denorm_inv4_len5_5
% SZS output start Proof for cubes_denorm_inv4_len5_5
1. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $difference($sum($product($uminus(6),1),X2),$product(6,X3)) & 0 = $sum($sum($difference($difference($difference($product(4,1),$product(18,X0)),$product($product(9,X0),X3)),$product(5,X1)),$product(9,X3)),$product(X1,X1)) & 0 = $sum($difference($sum($difference(1,$product(3,X0)),$product(X1,X3)),X1),$product(2,X3)) & 0 = $sum($difference($sum(1,$product($product(3,X3),X3)),X1),$product(3,X3)) & ~0 = $product($sum($difference($sum(1,$product($product(3,$sum(X3,1)),$sum(X3,1))),$sum(X1,X2)),$product(3,$sum(X3,1))),$difference($sum($sum($sum($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X0,X1)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1))),$product($product($product($product($sum(X2,6),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$product($product($product($product($product($sum(X0,X1),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$product($product($product($product($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$sum(X3,1))))) [input]
2. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X3))) & 0 = $sum($sum($sum($difference($difference($product(4,1),$product(18,X0)),$product($product(9,X0),X3)),$uminus($product(5,X1))),$product(9,X3)),$product(X1,X1)) & 0 = $sum($sum($sum($difference(1,$product(3,X0)),$product(X1,X3)),$uminus(X1)),$product(2,X3)) & 0 = $sum($sum($sum(1,$product($product(3,X3),X3)),$uminus(X1)),$product(3,X3)) & ~0 = $product($sum($sum($sum(1,$product($product(3,$sum(X3,1)),$sum(X3,1))),$uminus($sum(X1,X2))),$product(3,$sum(X3,1))),$sum($sum($sum($sum($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X0,X1)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1))),$product($product($product($product($sum(X2,6),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$product($product($product($product($product($sum(X0,X1),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$uminus($product($product($product($product($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$sum(X3,1)))))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 141]
4. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 142]
5. $sum(X0,0) = X0 [theory axiom 143]
6. $uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)) [theory axiom 145]
7. 0 = $sum(X0,$uminus(X0)) [theory axiom 146]
8. ~$less(X0,X0) [theory axiom 148]
9. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) [theory axiom 149]
10. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom 150]
11. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom 151]
12. $less(X1,$sum(X0,1)) | $less(X0,X1) [theory axiom 153]
13. $uminus($uminus(X0)) = X0 [theory axiom 154]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 141]
17. 0 = $product(X0,0) [theory axiom 155]
19. 0 = X0 | $product(X0,X2) != X1 | $product(X0,X3) != X1 | X2 = X3 [theory axiom 157]
20. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom 167]
21. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X3))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X0)),$uminus($product($product(9,X0),X3))),$uminus($product(5,X1))),$product(9,X3)),$product(X1,X1)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X0))),$product(X1,X3)),$uminus(X1)),$product(2,X3)) & 0 = $sum($sum($sum(1,$product($product(3,X3),X3)),$uminus(X1)),$product(3,X3)) & ~0 = $product($sum($sum($sum(1,$product($product(3,$sum(X3,1)),$sum(X3,1))),$uminus($sum(X1,X2))),$product(3,$sum(X3,1))),$sum($sum($sum($sum($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X0,X1)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1))),$product($product($product($product($sum(X2,6),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$product($product($product($product($product($sum(X0,X1),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$uminus($product($product($product($product($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$sum(X3,1)))))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X1))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X3)),$uminus($product($product(9,X3),X1))),$uminus($product(5,X0))),$product(9,X1)),$product(X0,X0)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X3))),$product(X0,X1)),$uminus(X0)),$product(2,X1)) & 0 = $sum($sum($sum(1,$product($product(3,X1),X1)),$uminus(X0)),$product(3,X1)) & ~0 = $product($sum($sum($sum(1,$product($product(3,$sum(X1,1)),$sum(X1,1))),$uminus($sum(X0,X2))),$product(3,$sum(X1,1))),$sum($sum($sum($sum($product($product($product($product($sum(X0,X2),$sum(X0,X2)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$product($product($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X3,X0)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1))),$product($product($product($product($sum(X2,6),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1))),$product($product($product($product($product($sum(X3,X0),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1))),$uminus($product($product($product($product($product($product($product($product($sum(X0,X2),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$sum(X1,1)))))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X1))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X3)),$uminus($product($product(9,X3),X1))),$uminus($product(5,X0))),$product(9,X1)),$product(X0,X0)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X3))),$product(X0,X1)),$uminus(X0)),$product(2,X1)) & 0 = $sum($sum($sum(1,$product($product(3,X1),X1)),$uminus(X0)),$product(3,X1)) & 0 != $product($sum($sum($sum(1,$product($product(3,$sum(X1,1)),$sum(X1,1))),$uminus($sum(X0,X2))),$product(3,$sum(X1,1))),$sum($sum($sum($sum($product($product($product($product($sum(X0,X2),$sum(X0,X2)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$product($product($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X3,X0)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1))),$product($product($product($product($sum(X2,6),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1))),$product($product($product($product($product($sum(X3,X0),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1))),$uminus($product($product($product($product($product($product($product($product($sum(X0,X2),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$sum(X1,1)))))) [flattening 22]
24. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($sum($sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1))),$product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$uminus($product($product($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1))))) [cnf transformation 23]
25. 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) [cnf transformation 23]
26. 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) [cnf transformation 23]
27. 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [cnf transformation 23]
28. 0 = $sum($sum($product($uminus(6),1),sK2),$uminus($product(6,sK1))) [cnf transformation 23]
29. $product(X0,X2) != $product(X0,X3) | 0 = X0 | X2 = X3 [equality resolution 19]
30. 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [evaluation 27]
31. 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) [evaluation 28]
33. 1 <=> 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) [avatar definition]
35. 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) <- (1) [avatar component clause 33]
36. 1 [avatar split clause 31,33]
38. 2 <=> 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) [avatar definition]
40. 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) <- (2) [avatar component clause 38]
41. 2 [avatar split clause 25,38]
43. 3 <=> 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) [avatar definition]
45. 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) <- (3) [avatar component clause 43]
46. 3 [avatar split clause 26,43]
48. 4 <=> 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [avatar definition]
50. 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) <- (4) [avatar component clause 48]
51. 4 [avatar split clause 30,48]
52. 0 = $sum($uminus(X0),X0) [superposition 7,13]
56. 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) <- (2) [backward demodulation 40,3]
57. 0 = $sum($sum($uminus(sK0),$sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1))),$product(2,sK1)) <- (3) [backward demodulation 45,3]
58. 0 = $sum($sum($sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))),$product(9,sK1)),$product(sK0,sK0)) <- (4) [backward demodulation 50,3]
59. $sum(0,X0) = X0 [superposition 3,5]
67. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) <- (3) [forward demodulation 57,3]
68. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) <- (4) [forward demodulation 58,3]
70. 5 <=> 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) [avatar definition]
72. 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) <- (5) [avatar component clause 70]
73. 5 | ~2 [avatar split clause 56,38,70]
74. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 72,3]
75. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 72,3]
76. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 3,72]
77. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 3,72]
79. 6 <=> 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) [avatar definition]
81. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (6) [avatar component clause 79]
82. 6 | ~5 [avatar split clause 74,70,79]
83. 6 | ~5 [avatar split clause 75,70,79]
84. 6 | ~5 [avatar split clause 76,70,79]
85. 6 | ~5 [avatar split clause 77,70,79]
87. 7 <=> 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) [avatar definition]
89. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) <- (7) [avatar component clause 87]
90. 7 | ~3 [avatar split clause 67,43,87]
91. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 89,3]
92. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 89,3]
93. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 3,89]
94. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 3,89]
96. 8 <=> 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) [avatar definition]
99. 8 | ~7 [avatar split clause 91,87,96]
100. 8 | ~7 [avatar split clause 92,87,96]
101. 8 | ~7 [avatar split clause 93,87,96]
102. 8 | ~7 [avatar split clause 94,87,96]
104. 9 <=> 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) [avatar definition]
106. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) <- (9) [avatar component clause 104]
107. 9 | ~4 [avatar split clause 68,48,104]
108. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 106,3]
109. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 106,3]
110. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 3,106]
111. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 3,106]
113. 10 <=> 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) [avatar definition]
115. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (10) [avatar component clause 113]
116. 10 | ~9 [avatar split clause 108,104,113]
117. 10 | ~9 [avatar split clause 109,104,113]
118. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (6) [backward demodulation 81,14]
119. 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) <- (5) [backward demodulation 72,14]
120. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) <- (9) [backward demodulation 106,14]
122. 0 = $product(0,X1) [superposition 14,17]
130. 11 <=> 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) [avatar definition]
132. 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) <- (11) [avatar component clause 130]
133. 11 | ~5 [avatar split clause 119,70,130]
134. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 132,3]
135. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 132,3]
136. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 3,132]
137. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 3,132]
139. 12 <=> 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) [avatar definition]
142. 12 | ~6 [avatar split clause 118,79,139]
143. 12 | ~11 [avatar split clause 134,130,139]
144. 12 | ~11 [avatar split clause 135,130,139]
145. 12 | ~11 [avatar split clause 136,130,139]
146. 12 | ~11 [avatar split clause 137,130,139]
148. 13 <=> 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) [avatar definition]
150. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) <- (13) [avatar component clause 148]
151. 13 | ~9 [avatar split clause 120,104,148]
152. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 150,3]
153. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 150,3]
154. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 3,150]
155. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 3,150]
156. $less(X0,$sum(X0,1)) [resolution 12,8]
159. $less(X4,0) | $less($uminus(1),X4) [superposition 12,52]
160. $less(X4,0) | $less(-1,X4) [evaluation 159]
161. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (9) [forward demodulation 110,14]
162. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (9) [forward demodulation 111,14]
163. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (10) [forward demodulation 115,14]
165. 14 <=> 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) [avatar definition]
168. 14 | ~13 [avatar split clause 152,148,165]
169. 14 | ~13 [avatar split clause 153,148,165]
170. 14 | ~13 [avatar split clause 154,148,165]
171. 14 | ~13 [avatar split clause 155,148,165]
172. 14 | ~9 [avatar split clause 161,104,165]
173. 14 | ~9 [avatar split clause 162,104,165]
174. 14 | ~10 [avatar split clause 163,113,165]
181. $less(X2,1) | $less(0,X2) [superposition 12,59]
192. $less(X0,$sum(X1,1)) | ~$less(X0,X1) [resolution 156,9]
194. $less(X0,$sum(1,X0)) [superposition 156,3]
199. ~$less($sum(1,X2),$sum(X2,1)) [resolution 194,20]
245. ~$less($sum(X0,1),X0) [resolution 192,8]
248. ~$less($sum(1,X6),X6) [resolution 192,199]
258. ~$less($sum($sum(X2,1),1),X2) [resolution 245,192]
267. ~$less($sum(2,X2),X2) [evaluation 258]
292. ~$less($sum(2,$sum(X2,1)),X2) [resolution 267,192]
297. ~$less($sum(X0,2),X0) [superposition 267,3]
301. ~$less($sum(3,X2),X2) [evaluation 292]
349. ~$less($sum(3,$sum(X3,1)),X3) [resolution 301,192]
353. ~$less($sum(X0,3),X0) [superposition 301,3]
356. ~$less($sum(4,X3),X3) [evaluation 349]
421. $sum($sum(-6,sK2),$sum($uminus($product(6,sK1)),X12)) = $sum(0,X12) <- (1) [superposition 4,35]
431. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (1) [superposition 4,35]
439. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (1) [superposition 35,4]
457. $sum($sum(-6,sK2),$sum($uminus($product(6,sK1)),X12)) = X12 <- (1) [forward demodulation 421,59]
492. 15 <=> 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) [avatar definition]
494. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (15) [avatar component clause 492]
495. 15 | ~1 [avatar split clause 431,33,492]
496. 15 | ~1 [avatar split clause 439,33,492]
500. $product(6,sK1) = $sum($sum(-6,sK2),0) <- (1) [superposition 457,52]
503. $sum($sum(-6,sK2),0) = $uminus($uminus($product(6,sK1))) <- (1) [superposition 457,7]
509. $product(6,sK1) = $sum(-6,sK2) <- (1) [evaluation 503]
510. $product(6,sK1) = $sum(-6,sK2) <- (1) [evaluation 500]
513. 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) <- (1, 15) [backward demodulation 494,509]
516. 16 <=> $product(6,sK1) = $sum(-6,sK2) [avatar definition]
518. $product(6,sK1) = $sum(-6,sK2) <- (16) [avatar component clause 516]
519. 16 | ~1 [avatar split clause 509,33,516]
520. 16 | ~1 [avatar split clause 510,33,516]
522. 17 <=> 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) [avatar definition]
524. 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) <- (17) [avatar component clause 522]
525. 17 | ~1 | ~15 [avatar split clause 513,492,33,522]
526. $sum(0,X0) = $sum(-6,$sum($sum(sK2,$uminus($sum(-6,sK2))),X0)) <- (17) [superposition 4,524]
527. $sum(0,X0) = $sum(-6,$sum(sK2,$sum($uminus($sum(-6,sK2)),X0))) <- (17) [forward demodulation 526,4]
528. $sum(-6,$sum(sK2,$sum($uminus($sum(-6,sK2)),X0))) = X0 <- (17) [forward demodulation 527,59]
761. $sum(-6,sK2) != $product(6,X22) | 0 = 6 | sK1 = X22 <- (16) [superposition 29,518]
772. $sum(-6,sK2) != $product(6,X22) | sK1 = X22 <- (16) [evaluation 761]
787. 0 != $sum(-6,sK2) | 0 = sK1 <- (16) [superposition 772,17]
790. 18 <=> 0 = sK1 [avatar definition]
792. 0 = sK1 <- (18) [avatar component clause 790]
794. 19 <=> 0 = $sum(-6,sK2) [avatar definition]
797. 18 | ~19 | ~16 [avatar split clause 787,516,794,790]
798. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,0),$sum(1,$uminus($product(3,sK3))))),$product(2,0)) <- (7, 18) [backward demodulation 89,792]
800. 0 = $sum($sum($uminus(sK0),$sum(1,$product(0,$product(3,0)))),$product(3,0)) <- (11, 18) [backward demodulation 132,792]
802. 0 = $sum($sum($product(9,0),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(0,$product(9,sK3)))))),$product(sK0,sK0)) <- (13, 18) [backward demodulation 150,792]
804. $sum(-6,sK2) = $product(6,0) <- (16, 18) [backward demodulation 518,792]
815. 0 = $sum(-6,sK2) <- (16, 18) [evaluation 804]
817. 0 = $sum($sum(0,$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(0,sK3))))),$product(sK0,sK0)) <- (13, 18) [evaluation 802]
819. 0 = $sum(1,$uminus(sK0)) <- (11, 18) [evaluation 800]
821. 0 = $sum(1,$sum($uminus(sK0),$sum($product(sK0,0),$uminus($product(3,sK3))))) <- (7, 18) [evaluation 798]
822. 0 = $sum(1,$sum($uminus(sK0),$sum(0,$uminus($product(3,sK3))))) <- (7, 18) [forward demodulation 821,17]
823. 0 = $sum(1,$sum($uminus(sK0),$uminus($product(3,sK3)))) <- (7, 18) [evaluation 822]
824. 0 = $sum(1,$uminus($sum($product(3,sK3),sK0))) <- (7, 18) [forward demodulation 823,6]
825. 0 = $sum(1,$uminus($sum(sK0,$product(3,sK3)))) <- (7, 18) [forward demodulation 824,3]
830. 0 = $sum($sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(0,sK3)))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 817,59]
831. 0 = $sum($sum($uminus($product(5,sK0)),$sum($uminus($product(0,sK3)),$difference(4,$product(18,sK3)))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 830,3]
832. 0 = $sum($sum($uminus($product(5,sK0)),$sum($uminus(0),$difference(4,$product(18,sK3)))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 831,122]
833. 0 = $sum($sum($uminus($product(5,sK0)),$sum(0,$difference(4,$product(18,sK3)))),$product(sK0,sK0)) <- (13, 18) [evaluation 832]
834. 0 = $sum($sum($uminus($product(5,sK0)),$difference(4,$product(18,sK3))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 833,59]
845. 0 = $sum(-6,$sum(sK2,$uminus(0))) <- (16, 17, 18) [backward demodulation 524,815]
846. $sum(-6,$sum(sK2,$sum($uminus(0),X0))) = X0 <- (16, 17, 18) [backward demodulation 528,815]
848. 19 | ~16 | ~18 [avatar split clause 815,790,516,794]
849. $sum(-6,$sum(sK2,X0)) = X0 <- (16, 17, 18) [evaluation 846]
850. 0 = $sum(-6,sK2) <- (16, 17, 18) [evaluation 845]
851. 19 | ~16 | ~17 | ~18 [avatar split clause 850,790,522,516,794]
855. 20 <=> 0 = $sum(1,$uminus(sK0)) [avatar definition]
857. 0 = $sum(1,$uminus(sK0)) <- (20) [avatar component clause 855]
858. 20 | ~11 | ~18 [avatar split clause 819,790,130,855]
859. ~$less(0,$uminus(sK0)) <- (20) [superposition 248,857]
862. $less($uminus(sK0),0) <- (20) [superposition 194,857]
868. 21 <=> $less($uminus(sK0),0) [avatar definition]
870. $less($uminus(sK0),0) <- (21) [avatar component clause 868]
871. 21 | ~20 [avatar split clause 862,855,868]
872. $less($sum($uminus(sK0),X0),$sum(0,X0)) <- (21) [resolution 870,11]
874. ~$less(0,$sum($uminus(sK0),1)) <- (21) [resolution 870,20]
875. $less($sum($uminus(sK0),X0),X0) <- (21) [forward demodulation 872,59]
876. ~$less(0,$sum(1,$uminus(sK0))) <- (21) [forward demodulation 874,3]
879. 22 <=> 0 = $sum(1,$uminus($sum(sK0,$product(3,sK3)))) [avatar definition]
881. 0 = $sum(1,$uminus($sum(sK0,$product(3,sK3)))) <- (22) [avatar component clause 879]
882. 22 | ~7 | ~18 [avatar split clause 825,790,87,879]
884. 23 <=> $less(0,$uminus(sK0)) [avatar definition]
887. ~23 | ~20 [avatar split clause 859,855,884]
894. $less(0,sK0) <- (21) [superposition 875,52]
895. $less($uminus($sum(X0,sK0)),$uminus(X0)) <- (21) [superposition 875,6]
897. $less(0,$uminus($uminus(sK0))) <- (21) [superposition 875,7]
898. $less($sum(X1,$uminus(sK0)),X1) <- (21) [superposition 875,3]
900. $less(0,sK0) <- (21) [evaluation 897]
904. 24 <=> $less(0,sK0) [avatar definition]
906. $less(0,sK0) <- (24) [avatar component clause 904]
907. 24 | ~21 [avatar split clause 894,868,904]
908. 24 | ~21 [avatar split clause 900,868,904]
910. ~$less(X1,0) | $less(X1,sK0) <- (24) [resolution 906,9]
911. ~$less(sK0,$sum(0,1)) <- (24) [resolution 906,20]
912. ~$less(sK0,1) <- (24) [evaluation 911]
915. 25 <=> $less(sK0,1) [avatar definition]
917. ~$less(sK0,1) <- (~25) [avatar component clause 915]
918. ~25 | ~24 [avatar split clause 912,904,915]
920. $less(1,sK0) | 1 = sK0 <- (~25) [resolution 917,10]
964. $less(X0,sK0) | $less(-1,X0) <- (24) [resolution 910,160]
968. $less($sum($uminus(sK0),0),sK0) <- (21, 24) [resolution 910,875]
969. $less($sum(0,$uminus(sK0)),sK0) <- (21, 24) [resolution 910,898]
970. $less($uminus(sK0),sK0) <- (21, 24) [resolution 910,870]
971. $less(-1,sK0) | $less(0,0) <- (24) [resolution 910,160]
972. $less(-1,sK0) <- (24) [evaluation 971]
973. $less($uminus(sK0),sK0) <- (21, 24) [forward demodulation 968,5]
974. $less($uminus(sK0),sK0) <- (21, 24) [forward demodulation 969,59]
976. 26 <=> $less(-1,sK0) [avatar definition]
978. $less(-1,sK0) <- (26) [avatar component clause 976]
979. 26 | ~24 [avatar split clause 972,904,976]
981. ~$less(X1,-1) | $less(X1,sK0) <- (26) [resolution 978,9]
982. ~$less(sK0,$sum(-1,1)) <- (26) [resolution 978,20]
983. ~$less(sK0,0) <- (26) [evaluation 982]
985. 27 <=> $less(sK0,0) [avatar definition]
988. ~27 | ~26 [avatar split clause 983,976,985]
993. 28 <=> $less($uminus(sK0),sK0) [avatar definition]
995. $less($uminus(sK0),sK0) <- (28) [avatar component clause 993]
996. 28 | ~21 | ~24 [avatar split clause 970,904,868,993]
998. 29 <=> 0 = $sum($sum($uminus($product(5,sK0)),$difference(4,$product(18,sK3))),$product(sK0,sK0)) [avatar definition]
1000. 0 = $sum($sum($uminus($product(5,sK0)),$difference(4,$product(18,sK3))),$product(sK0,sK0)) <- (29) [avatar component clause 998]
1001. 29 | ~13 | ~18 [avatar split clause 834,790,148,998]
1002. 28 | ~21 | ~24 [avatar split clause 973,904,868,993]
1003. 28 | ~21 | ~24 [avatar split clause 974,904,868,993]
1006. ~$less(sK0,$sum($uminus(sK0),1)) <- (28) [resolution 995,20]
1007. ~$less(sK0,$sum(1,$uminus(sK0))) <- (28) [forward demodulation 1006,3]
1010. 30 <=> 1 = sK0 [avatar definition]
1012. 1 = sK0 <- (30) [avatar component clause 1010]
1014. 31 <=> $less(1,sK0) [avatar definition]
1016. $less(1,sK0) <- (31) [avatar component clause 1014]
1017. 30 | 31 | 25 [avatar split clause 920,915,1014,1010]
1021. 0 = $sum(1,$uminus($sum(1,$product(3,sK3)))) <- (22, 30) [backward demodulation 881,1012]
1025. ~$less(X1,0) | $less(X1,1) <- (24, 30) [backward demodulation 910,1012]
1055. 32 <=> 0 = $sum(1,$uminus($sum(1,$product(3,sK3)))) [avatar definition]
1057. 0 = $sum(1,$uminus($sum(1,$product(3,sK3)))) <- (32) [avatar component clause 1055]
1058. 32 | ~22 | ~30 [avatar split clause 1021,1010,879,1055]
1059. $less(X0,1) | $less(-1,X0) <- (24, 30) [resolution 1025,160]
1113. $less($uminus($sum(X0,1)),$uminus(X0)) <- (21, 30) [forward demodulation 895,1012]
1117. $less($uminus($sum(1,X0)),$uminus(X0)) <- (21, 30) [superposition 1113,3]
1122. $less($uminus($sum($uminus(X0),1)),X0) <- (21, 30) [superposition 1113,13]
1124. $less($uminus($sum(1,$uminus(X0))),X0) <- (21, 30) [forward demodulation 1122,3]
1129. $less($uminus(0),$uminus($uminus($sum(1,$product(3,sK3))))) <- (21, 30, 32) [superposition 1117,1057]
1135. $less(0,$sum(1,$product(3,sK3))) <- (21, 30, 32) [evaluation 1129]
1144. $less($uminus(0),$sum(1,$product(3,sK3))) <- (21, 30, 32) [superposition 1124,1057]
1146. $less(0,$sum(1,$product(3,sK3))) <- (21, 30, 32) [evaluation 1144]
1150. 33 <=> $less(0,$sum(1,$product(3,sK3))) [avatar definition]
1152. $less(0,$sum(1,$product(3,sK3))) <- (33) [avatar component clause 1150]
1153. 33 | ~21 | ~30 | ~32 [avatar split clause 1135,1055,1010,868,1150]
1154. 0 = $sum($sum($uminus($product(5,1)),$difference(4,$product(18,sK3))),$product(1,1)) <- (29, 30) [forward demodulation 1000,1012]
1155. 0 = $sum(-4,$difference(4,$product(18,sK3))) <- (29, 30) [evaluation 1154]
1156. 33 | ~21 | ~30 | ~32 [avatar split clause 1146,1055,1010,868,1150]
1159. ~$less($sum(1,$product(3,sK3)),$sum(0,1)) <- (33) [resolution 1152,20]
1160. ~$less($sum(1,$product(3,sK3)),1) <- (33) [evaluation 1159]
1169. 34 <=> $less($sum(1,$product(3,sK3)),1) [avatar definition]
1171. ~$less($sum(1,$product(3,sK3)),1) <- (~34) [avatar component clause 1169]
1172. ~34 | ~33 [avatar split clause 1160,1150,1169]
1188. $less(-1,$sum(1,$product(3,sK3))) <- (24, 30, ~34) [resolution 1171,1059]
1210. 35 <=> 0 = $sum(-4,$difference(4,$product(18,sK3))) [avatar definition]
1213. 35 | ~29 | ~30 [avatar split clause 1155,1010,998,1210]
1230. 36 <=> $less(-1,$sum(1,$product(3,sK3))) [avatar definition]
1232. $less(-1,$sum(1,$product(3,sK3))) <- (36) [avatar component clause 1230]
1233. 36 | ~24 | ~30 | 34 [avatar split clause 1188,1169,1010,904,1230]
1236. ~$less($sum(1,$product(3,sK3)),$sum(-1,1)) <- (36) [resolution 1232,20]
1237. ~$less($sum(1,$product(3,sK3)),0) <- (36) [evaluation 1236]
1258. 37 <=> $less($sum(1,$product(3,sK3)),0) [avatar definition]
1261. ~37 | ~36 [avatar split clause 1237,1230,1258]
1268. $uminus(sK2) = $sum(-6,0) <- (16, 17, 18) [superposition 849,7]
1272. -6 = $uminus(sK2) <- (16, 17, 18) [evaluation 1268]
1275. 38 <=> -6 = $uminus(sK2) [avatar definition]
1277. -6 = $uminus(sK2) <- (38) [avatar component clause 1275]
1278. 38 | ~16 | ~17 | ~18 [avatar split clause 1272,790,522,516,1275]
1282. sK2 = $uminus(-6) <- (38) [superposition 13,1277]
1288. 6 = sK2 <- (38) [evaluation 1282]
1295. 39 <=> 6 = sK2 [avatar definition]
1297. 6 = sK2 <- (39) [avatar component clause 1295]
1298. 39 | ~38 [avatar split clause 1288,1275,1295]
1305. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1))),$product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($product($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)))))) [forward demodulation 24,4]
1306. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1))),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($product($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1))))))) [forward demodulation 1305,4]
1307. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($product($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)))))))) [forward demodulation 1306,4]
1308. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))))))))) [forward demodulation 1307,14]
1309. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)))))))))) [forward demodulation 1308,14]
1310. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))))))))))) [forward demodulation 1309,14]
1311. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK2,6),$product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)))))))))))) [forward demodulation 1310,14]
1312. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6))))))))))))) [forward demodulation 1311,14]
1313. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)))))))))))))) [forward demodulation 1312,14]
1314. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(sK2,6),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1313,14]
1315. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,6))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,6),$sum(sK0,6)),$sum(6,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,6)),$sum(6,6)),$sum(6,6)),$sum(6,6)),$sum(sK1,1)),$sum($product($product($product($product($sum(6,6),$sum(6,6)),$sum(6,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(6,6)),$sum(6,6)),$sum(6,6)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($sum(6,6),$product($sum(6,6),$product($sum(6,6),$product($sum(sK0,6),$product($sum(sK0,6),$sum(sK0,6))))))))))))))) <- (39) [forward demodulation 1314,1297]
1316. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,6))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,6),$sum(sK0,6)),12),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product(1728,$product($sum(sK3,sK0),$product($sum(sK3,sK0),$product($sum(sK3,sK0),$sum(sK0,6))))),$sum(sK1,1)),$sum($product($product(1728,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product(1728,$sum(sK3,sK0)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product(1728,$product($sum(sK0,6),$product($sum(sK0,6),$sum(sK0,6))))))))))))) <- (39) [evaluation 1315]
1317. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(6,sK0))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(6,sK0),$sum(6,sK0)),12),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product(1728,$product($sum(sK3,sK0),$product($sum(sK3,sK0),$product($sum(sK3,sK0),$sum(6,sK0))))),$sum(sK1,1)),$sum($product($product(1728,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product(1728,$sum(sK3,sK0)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product(1728,$product($sum(6,sK0),$product($sum(6,sK0),$sum(6,sK0))))))))))))) <- (39) [forward demodulation 1316,3]
1318. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(6,1))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(6,1),$sum(6,1)),12),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product(1728,$product($sum(sK3,1),$product($sum(sK3,1),$product($sum(sK3,1),$sum(6,1))))),$sum(sK1,1)),$sum($product($product(1728,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product(1728,$sum(sK3,1)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product(1728,$product($sum(6,1),$product($sum(6,1),$sum(6,1))))))))))))) <- (30, 39) [forward demodulation 1317,1012]
1319. 0 != $product($sum($sum(-6,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$product(3,$sum(sK1,1))),$sum($product($product(588,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product(12096,$product($sum(sK3,1),$product($sum(sK3,1),$sum(sK3,1)))),$sum(sK1,1)),$sum($product($product(1728,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product(1728,$sum(sK3,1)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),592704))))))))) <- (30, 39) [evaluation 1318]
1320. 0 != $product($sum($sum(-6,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$product(3,$sum(sK1,1))),$sum($product($product(588,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product(12096,$product($sum(sK3,1),$product($sum(sK3,1),$sum(sK3,1)))),$sum(sK1,1)),$sum($product($product(1728,$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product(1728,$sum(sK3,1)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product(592704,$sum(sK1,1)))))))))) <- (30, 39) [forward demodulation 1319,14]
1321. 0 != $product($sum($sum(-6,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$product(3,$sum(1,sK1))),$sum($product($product(588,$sum(1,sK1)),$sum(1,sK1)),$sum($product($product(12096,$product($sum(sK3,1),$product($sum(sK3,1),$sum(sK3,1)))),$sum(1,sK1)),$sum($product($product(1728,$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product(1728,$sum(sK3,1)),$sum(1,sK1)),$sum(1,sK1)),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product(592704,$sum(1,sK1)))))))))) <- (30, 39) [forward demodulation 1320,3]
1322. 0 != $product($sum($sum(-6,$product($product(3,$sum(1,0)),$sum(1,0))),$product(3,$sum(1,0))),$sum($product($product(588,$sum(1,0)),$sum(1,0)),$sum($product($product(12096,$product($sum(sK3,1),$product($sum(sK3,1),$sum(sK3,1)))),$sum(1,0)),$sum($product($product(1728,$sum(1,0)),$sum(1,0)),$sum($product($product($product(1728,$sum(sK3,1)),$sum(1,0)),$sum(1,0)),$uminus($product($sum(1,0),$product($sum(1,0),$product(592704,$sum(1,0)))))))))) <- (18, 30, 39) [forward demodulation 1321,792]
1323. 0 != $product(0,$sum(-590388,$sum($product(12096,$product($sum(sK3,1),$product($sum(sK3,1),$sum(sK3,1)))),$product(1728,$sum(sK3,1))))) <- (18, 30, 39) [evaluation 1322]
1324. $false <- (18, 30, 39) [subsumption resolution 1323,122]
1325. ~18 | ~30 | ~39 [avatar contradiction clause 1324]
1326. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(sK1,1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1314,3]
1327. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($sum(sK3,sK0),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1326,3]
1328. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($product($product($product($sum(sK3,sK0),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1327,14]
1329. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($product($product($sum(sK3,sK0),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1328,14]
1330. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($product($sum(sK3,sK0),$sum(6,sK2)),$sum(6,sK2))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1329,14]
1331. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK3,sK0),$sum(6,sK2)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1330,14]
1332. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK3,sK0)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1331,14]
1333. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1332,3]
1334. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1333,14]
1335. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($product($sum(6,sK2),$sum(6,sK2)),$sum(6,sK2)))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1334,14]
1336. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1335,14]
1337. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1336,14]
1338. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1337,14]
1339. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1338,14]
1340. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1339,14]
1341. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK0,sK3))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1340,14]
1342. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1341,14]
1343. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)),$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1342,14]
1344. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(6,sK2)))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1343,14]
1345. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1344,14]
1346. 0 != $product($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1)))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1345,4]
1347. 0 != $product($sum(1,$sum($product($product(3,$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1346,4]
1348. 0 != $product($sum(1,$sum($product($sum(1,sK1),$product(3,$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [forward demodulation 1347,14]
1362. ~$less(X1,1) | $less(X1,sK0) <- (31) [resolution 1016,9]
1363. ~$less(sK0,$sum(1,1)) <- (31) [resolution 1016,20]
1364. ~$less(sK0,2) <- (31) [evaluation 1363]
1369. 40 <=> $less(sK0,2) [avatar definition]
1371. ~$less(sK0,2) <- (~40) [avatar component clause 1369]
1372. ~40 | ~31 [avatar split clause 1364,1014,1369]
1373. $less(2,sK0) | 2 = sK0 <- (~40) [resolution 1371,10]
1377. ~$less(0,$sum($uminus(sK0),1)) <- (21) [resolution 870,20]
1379. ~$less(0,$sum(1,$uminus(sK0))) <- (21) [forward demodulation 1377,3]
1385. ~$less(sK0,$sum($uminus(sK0),1)) <- (28) [resolution 995,20]
1386. ~$less(sK0,$sum(1,$uminus(sK0))) <- (28) [forward demodulation 1385,3]
1417. 41 <=> $less(0,$sum(1,$uminus(sK0))) [avatar definition]
1420. ~41 | ~21 [avatar split clause 876,868,1417]
1421. ~41 | ~21 [avatar split clause 1379,868,1417]
1455. $less(-1,$sum(sK0,1)) <- (24) [resolution 964,245]
1456. $less(-1,$sum(sK0,2)) <- (24) [resolution 964,297]
1457. $less(-1,$sum(sK0,3)) <- (24) [resolution 964,353]
1458. $less(-1,$sum(1,sK0)) <- (24) [resolution 964,248]
1459. $less(-1,$sum(3,sK0)) <- (24) [resolution 964,301]
1460. $less(-1,$sum(2,sK0)) <- (24) [resolution 964,267]
1461. $less(-1,$sum(4,sK0)) <- (24) [resolution 964,356]
1468. $less(-1,$sum(1,sK0)) <- (24) [forward demodulation 1455,3]
1469. $less(-1,$sum(2,sK0)) <- (24) [forward demodulation 1456,3]
1470. $less(-1,$sum(3,sK0)) <- (24) [forward demodulation 1457,3]
1472. 42 <=> $less(-1,$sum(1,sK0)) [avatar definition]
1474. $less(-1,$sum(1,sK0)) <- (42) [avatar component clause 1472]
1475. 42 | ~24 [avatar split clause 1458,904,1472]
1477. 43 <=> $less(-1,$sum(3,sK0)) [avatar definition]
1479. $less(-1,$sum(3,sK0)) <- (43) [avatar component clause 1477]
1480. 43 | ~24 [avatar split clause 1459,904,1477]
1482. 44 <=> $less(-1,$sum(2,sK0)) [avatar definition]
1484. $less(-1,$sum(2,sK0)) <- (44) [avatar component clause 1482]
1485. 44 | ~24 [avatar split clause 1460,904,1482]
1486. 42 | ~24 [avatar split clause 1468,904,1472]
1487. 44 | ~24 [avatar split clause 1469,904,1482]
1488. 43 | ~24 [avatar split clause 1470,904,1477]
1491. ~$less($sum(1,sK0),$sum(-1,1)) <- (42) [resolution 1474,20]
1492. ~$less($sum(1,sK0),0) <- (42) [evaluation 1491]
1496. ~$less($sum(3,sK0),$sum(-1,1)) <- (43) [resolution 1479,20]
1497. ~$less($sum(3,sK0),0) <- (43) [evaluation 1496]
1501. ~$less($sum(2,sK0),$sum(-1,1)) <- (44) [resolution 1484,20]
1502. ~$less($sum(2,sK0),0) <- (44) [evaluation 1501]
1505. 45 <=> $less($sum(1,sK0),0) [avatar definition]
1508. ~45 | ~42 [avatar split clause 1492,1472,1505]
1510. 46 <=> $less($sum(3,sK0),0) [avatar definition]
1513. ~46 | ~43 [avatar split clause 1497,1477,1510]
1515. 47 <=> $less($sum(2,sK0),0) [avatar definition]
1518. ~47 | ~44 [avatar split clause 1502,1482,1515]
1529. 48 <=> $less(-1,$sum(4,sK0)) [avatar definition]
1531. $less(-1,$sum(4,sK0)) <- (48) [avatar component clause 1529]
1532. 48 | ~24 [avatar split clause 1461,904,1529]
1547. ~$less($sum(4,sK0),$sum(-1,1)) <- (48) [resolution 1531,20]
1548. ~$less($sum(4,sK0),0) <- (48) [evaluation 1547]
1551. 49 <=> $less($sum(4,sK0),0) [avatar definition]
1554. ~49 | ~48 [avatar split clause 1548,1529,1551]
1561. $less($sum($uminus(sK0),-1),sK0) <- (21, 26) [resolution 981,875]
1562. $less($sum(-1,$uminus(sK0)),sK0) <- (21, 26) [resolution 981,898]
1565. $less($sum(-1,$uminus(sK0)),sK0) <- (21, 26) [forward demodulation 1561,3]
1567. 50 <=> $less(sK0,$sum(1,$uminus(sK0))) [avatar definition]
1570. ~50 | ~28 [avatar split clause 1007,993,1567]
1571. $less(X0,sK0) | $less(0,X0) <- (31) [resolution 1362,181]
1585. ~50 | ~28 [avatar split clause 1386,993,1567]
1590. 51 <=> 2 = sK0 [avatar definition]
1594. 52 <=> $less(2,sK0) [avatar definition]
1596. $less(2,sK0) <- (52) [avatar component clause 1594]
1597. 51 | 52 | 40 [avatar split clause 1373,1369,1594,1590]
1600. ~$less(sK0,$sum(2,1)) <- (52) [resolution 1596,20]
1601. ~$less(sK0,3) <- (52) [evaluation 1600]
1603. 53 <=> $less(sK0,3) [avatar definition]
1606. ~53 | ~52 [avatar split clause 1601,1594,1603]
1610. 54 <=> $less($sum(-1,$uminus(sK0)),sK0) [avatar definition]
1613. 54 | ~21 | ~26 [avatar split clause 1562,976,868,1610]
1614. 54 | ~21 | ~26 [avatar split clause 1565,976,868,1610]
1633. $less(0,$sum(sK0,1)) <- (31) [resolution 1571,245]
1634. $less(0,$sum(sK0,2)) <- (31) [resolution 1571,297]
1635. $less(0,$sum(sK0,3)) <- (31) [resolution 1571,353]
1636. $less(0,$sum(1,sK0)) <- (31) [resolution 1571,248]
1637. $less(0,$sum(3,sK0)) <- (31) [resolution 1571,301]
1638. $less(0,$sum(2,sK0)) <- (31) [resolution 1571,267]
1639. $less(0,$sum(4,sK0)) <- (31) [resolution 1571,356]
1650. $less(0,$sum(1,sK0)) <- (31) [forward demodulation 1633,3]
1651. $less(0,$sum(2,sK0)) <- (31) [forward demodulation 1634,3]
1652. $less(0,$sum(3,sK0)) <- (31) [forward demodulation 1635,3]
1655. 55 <=> $less(0,$sum(1,sK0)) [avatar definition]
1657. $less(0,$sum(1,sK0)) <- (55) [avatar component clause 1655]
1658. 55 | ~31 [avatar split clause 1636,1014,1655]
1660. 56 <=> $less(0,$sum(3,sK0)) [avatar definition]
1662. $less(0,$sum(3,sK0)) <- (56) [avatar component clause 1660]
1663. 56 | ~31 [avatar split clause 1637,1014,1660]
1665. 57 <=> $less(0,$sum(2,sK0)) [avatar definition]
1667. $less(0,$sum(2,sK0)) <- (57) [avatar component clause 1665]
1668. 57 | ~31 [avatar split clause 1638,1014,1665]
1669. 55 | ~31 [avatar split clause 1650,1014,1655]
1670. 57 | ~31 [avatar split clause 1651,1014,1665]
1671. 56 | ~31 [avatar split clause 1652,1014,1660]
1674. ~$less($sum(1,sK0),$sum(0,1)) <- (55) [resolution 1657,20]
1675. ~$less($sum(1,sK0),1) <- (55) [evaluation 1674]
1680. ~$less($sum(3,sK0),$sum(0,1)) <- (56) [resolution 1662,20]
1681. ~$less($sum(3,sK0),1) <- (56) [evaluation 1680]
1686. ~$less($sum(2,sK0),$sum(0,1)) <- (57) [resolution 1667,20]
1687. ~$less($sum(2,sK0),1) <- (57) [evaluation 1686]
1705. 58 <=> $less($sum(1,sK0),1) [avatar definition]
1708. ~58 | ~55 [avatar split clause 1675,1655,1705]
1710. 59 <=> $less($sum(3,sK0),1) [avatar definition]
1713. ~59 | ~56 [avatar split clause 1681,1660,1710]
1715. 60 <=> $less($sum(2,sK0),1) [avatar definition]
1718. ~60 | ~57 [avatar split clause 1687,1665,1715]
1729. 61 <=> $less(0,$sum(4,sK0)) [avatar definition]
1731. $less(0,$sum(4,sK0)) <- (61) [avatar component clause 1729]
1732. 61 | ~31 [avatar split clause 1639,1014,1729]
1735. ~$less($sum(4,sK0),$sum(0,1)) <- (61) [resolution 1731,20]
1736. ~$less($sum(4,sK0),1) <- (61) [evaluation 1735]
1740. 62 <=> $less($sum(4,sK0),1) [avatar definition]
1743. ~62 | ~61 [avatar split clause 1736,1729,1740]
1748. 63 <=> 0 = $product($sum(1,$sum($product($sum(1,sK1),$product(3,$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))),$sum($product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$sum(sK0,sK3)))))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))))))))))) [avatar definition]
1751. ~63 [avatar split clause 1348,1748]
1752. $false [avatar sat refutation 36,41,46,51,73,82,83,84,85,90,99,100,101,102,107,116,117,133,142,143,144,145,146,151,168,169,170,171,172,173,174,495,496,519,520,525,797,848,851,858,871,882,887,907,908,918,979,988,996,1001,1002,1003,1017,1058,1153,1156,1172,1213,1233,1261,1278,1298,1325,1372,1420,1421,1475,1480,1485,1486,1487,1488,1508,1513,1518,1532,1554,1570,1585,1597,1606,1613,1614,1658,1663,1668,1669,1670,1671,1708,1713,1718,1732,1743,1751]
% SZS output end Proof for cubes_denorm_inv4_len5_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Refutation

% Memory used [KB]: 8187
% Time elapsed: 0.270 s
% ------------------------------
% ------------------------------
% Success in time 1.775 s
