% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on cubes_denorm_inv3_len4_5
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 7931
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_6 on cubes_denorm_inv3_len4_5
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity:gtg=exists_all_2 on cubes_denorm_inv3_len4_5
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for cubes_denorm_inv3_len4_5
% SZS output start Proof for cubes_denorm_inv3_len4_5
1. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $difference($sum($product($uminus(6),1),X2),$product(6,X3)) & 0 = $sum($sum($difference($difference($difference($product(4,1),$product(18,X0)),$product($product(9,X0),X3)),$product(5,X1)),$product(9,X3)),$product(X1,X1)) & 0 = $sum($difference($sum($difference(1,$product(3,X0)),$product(X1,X3)),X1),$product(2,X3)) & 0 = $sum($difference($sum(1,$product($product(3,X3),X3)),X1),$product(3,X3)) & ~0 = $product($sum($difference($sum($difference(1,$product(3,$sum(X0,X1))),$product($sum(X1,X2),$sum(X3,1))),$sum(X1,X2)),$product(2,$sum(X3,1))),$sum($difference($difference($product($product($product($product($product($product($product($product($product($uminus(1),$sum(X0,X1)),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($sum(X0,X1),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6))),$product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$product($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1))))) [input]
2. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X3))) & 0 = $sum($sum($sum($difference($difference($product(4,1),$product(18,X0)),$product($product(9,X0),X3)),$uminus($product(5,X1))),$product(9,X3)),$product(X1,X1)) & 0 = $sum($sum($sum($difference(1,$product(3,X0)),$product(X1,X3)),$uminus(X1)),$product(2,X3)) & 0 = $sum($sum($sum(1,$product($product(3,X3),X3)),$uminus(X1)),$product(3,X3)) & ~0 = $product($sum($sum($sum($difference(1,$product(3,$sum(X0,X1))),$product($sum(X1,X2),$sum(X3,1))),$uminus($sum(X1,X2))),$product(2,$sum(X3,1))),$sum($sum($difference($product($product($product($product($product($product($product($product($product($uminus(1),$sum(X0,X1)),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($sum(X0,X1),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6))),$uminus($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)))),$product($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1))))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 141]
4. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 142]
5. $sum(X0,0) = X0 [theory axiom 143]
6. $uminus($sum(X0,X1)) = $sum($uminus(X1),$uminus(X0)) [theory axiom 145]
7. 0 = $sum(X0,$uminus(X0)) [theory axiom 146]
8. ~$less(X0,X0) [theory axiom 148]
9. ~$less(X1,X2) | ~$less(X0,X1) | $less(X0,X2) [theory axiom 149]
10. $less(X0,X1) | $less(X1,X0) | X0 = X1 [theory axiom 150]
11. ~$less(X0,X1) | $less($sum(X0,X2),$sum(X1,X2)) [theory axiom 151]
12. $less(X1,$sum(X0,1)) | $less(X0,X1) [theory axiom 153]
13. $uminus($uminus(X0)) = X0 [theory axiom 154]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 141]
17. 0 = $product(X0,0) [theory axiom 155]
19. 0 = X0 | $product(X0,X2) != X1 | $product(X0,X3) != X1 | X2 = X3 [theory axiom 157]
20. ~$less(X0,X1) | ~$less(X1,$sum(X0,1)) [theory axiom 167]
21. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X3))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X0)),$uminus($product($product(9,X0),X3))),$uminus($product(5,X1))),$product(9,X3)),$product(X1,X1)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X0))),$product(X1,X3)),$uminus(X1)),$product(2,X3)) & 0 = $sum($sum($sum(1,$product($product(3,X3),X3)),$uminus(X1)),$product(3,X3)) & ~0 = $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(X0,X1)))),$product($sum(X1,X2),$sum(X3,1))),$uminus($sum(X1,X2))),$product(2,$sum(X3,1))),$sum($sum($sum($product($product($product($product($product($product($product($product($product($uminus(1),$sum(X0,X1)),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$sum(X3,1)),$uminus($product($product($product($product($product($sum(X0,X1),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),$uminus($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)))),$product($product($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1))))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X1))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X3)),$uminus($product($product(9,X3),X1))),$uminus($product(5,X0))),$product(9,X1)),$product(X0,X0)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X3))),$product(X0,X1)),$uminus(X0)),$product(2,X1)) & 0 = $sum($sum($sum(1,$product($product(3,X1),X1)),$uminus(X0)),$product(3,X1)) & ~0 = $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(X3,X0)))),$product($sum(X0,X2),$sum(X1,1))),$uminus($sum(X0,X2))),$product(2,$sum(X1,1))),$sum($sum($sum($product($product($product($product($product($product($product($product($product($uminus(1),$sum(X3,X0)),$sum(X3,X0)),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$sum(X1,1)),$uminus($product($product($product($product($product($sum(X3,X0),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),$uminus($product($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)))),$product($product($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1))))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X1))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X3)),$uminus($product($product(9,X3),X1))),$uminus($product(5,X0))),$product(9,X1)),$product(X0,X0)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X3))),$product(X0,X1)),$uminus(X0)),$product(2,X1)) & 0 = $sum($sum($sum(1,$product($product(3,X1),X1)),$uminus(X0)),$product(3,X1)) & 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(X3,X0)))),$product($sum(X0,X2),$sum(X1,1))),$uminus($sum(X0,X2))),$product(2,$sum(X1,1))),$sum($sum($sum($product($product($product($product($product($product($product($product($product($uminus(1),$sum(X3,X0)),$sum(X3,X0)),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$sum(X1,1)),$uminus($product($product($product($product($product($sum(X3,X0),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),$uminus($product($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)))),$product($product($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1))))) [flattening 22]
24. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($sum($sum($product($product($product($product($product($product($product($product($product($uminus(1),$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))),$uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)))),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)))) [cnf transformation 23]
25. 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) [cnf transformation 23]
26. 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) [cnf transformation 23]
27. 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [cnf transformation 23]
28. 0 = $sum($sum($product($uminus(6),1),sK2),$uminus($product(6,sK1))) [cnf transformation 23]
29. $product(X0,X2) != $product(X0,X3) | 0 = X0 | X2 = X3 [equality resolution 19]
30. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($sum($sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))),$uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)))),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)))) [evaluation 24]
31. 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [evaluation 27]
32. 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) [evaluation 28]
34. 1 <=> 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) [avatar definition]
36. 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) <- (1) [avatar component clause 34]
37. 1 [avatar split clause 32,34]
39. 2 <=> 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) [avatar definition]
41. 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) <- (2) [avatar component clause 39]
42. 2 [avatar split clause 25,39]
44. 3 <=> 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) [avatar definition]
46. 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) <- (3) [avatar component clause 44]
47. 3 [avatar split clause 26,44]
49. 4 <=> 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [avatar definition]
51. 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) <- (4) [avatar component clause 49]
52. 4 [avatar split clause 31,49]
53. 0 = $sum($uminus(X0),X0) [superposition 7,13]
57. 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) <- (2) [backward demodulation 41,3]
58. 0 = $sum($sum($uminus(sK0),$sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1))),$product(2,sK1)) <- (3) [backward demodulation 46,3]
59. 0 = $sum($sum($sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))),$product(9,sK1)),$product(sK0,sK0)) <- (4) [backward demodulation 51,3]
60. $sum(0,X0) = X0 [superposition 3,5]
68. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) <- (3) [forward demodulation 58,3]
69. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) <- (4) [forward demodulation 59,3]
71. 5 <=> 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) [avatar definition]
73. 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) <- (5) [avatar component clause 71]
74. 5 | ~2 [avatar split clause 57,39,71]
75. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 73,3]
76. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 73,3]
77. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 3,73]
78. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 3,73]
80. 6 <=> 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) [avatar definition]
82. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (6) [avatar component clause 80]
83. 6 | ~5 [avatar split clause 75,71,80]
84. 6 | ~5 [avatar split clause 76,71,80]
85. 6 | ~5 [avatar split clause 77,71,80]
86. 6 | ~5 [avatar split clause 78,71,80]
88. 7 <=> 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) [avatar definition]
90. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) <- (7) [avatar component clause 88]
91. 7 | ~3 [avatar split clause 68,44,88]
92. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 90,3]
93. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 90,3]
94. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 3,90]
95. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 3,90]
97. 8 <=> 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) [avatar definition]
100. 8 | ~7 [avatar split clause 92,88,97]
101. 8 | ~7 [avatar split clause 93,88,97]
102. 8 | ~7 [avatar split clause 94,88,97]
103. 8 | ~7 [avatar split clause 95,88,97]
105. 9 <=> 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) [avatar definition]
107. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) <- (9) [avatar component clause 105]
108. 9 | ~4 [avatar split clause 69,49,105]
109. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 107,3]
110. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 107,3]
111. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 3,107]
112. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 3,107]
114. 10 <=> 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) [avatar definition]
116. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (10) [avatar component clause 114]
117. 10 | ~9 [avatar split clause 109,105,114]
118. 10 | ~9 [avatar split clause 110,105,114]
119. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (6) [backward demodulation 82,14]
120. 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) <- (5) [backward demodulation 73,14]
121. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) <- (9) [backward demodulation 107,14]
123. 0 = $product(0,X1) [superposition 14,17]
131. 11 <=> 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) [avatar definition]
133. 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) <- (11) [avatar component clause 131]
134. 11 | ~5 [avatar split clause 120,71,131]
135. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 133,3]
136. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 133,3]
137. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 3,133]
138. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 3,133]
140. 12 <=> 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) [avatar definition]
143. 12 | ~6 [avatar split clause 119,80,140]
144. 12 | ~11 [avatar split clause 135,131,140]
145. 12 | ~11 [avatar split clause 136,131,140]
146. 12 | ~11 [avatar split clause 137,131,140]
147. 12 | ~11 [avatar split clause 138,131,140]
149. 13 <=> 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) [avatar definition]
151. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) <- (13) [avatar component clause 149]
152. 13 | ~9 [avatar split clause 121,105,149]
153. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 151,3]
154. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 151,3]
155. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 3,151]
156. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 3,151]
157. $less(X0,$sum(X0,1)) [resolution 12,8]
160. $less(X4,0) | $less($uminus(1),X4) [superposition 12,53]
161. $less(X4,0) | $less(-1,X4) [evaluation 160]
162. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (9) [forward demodulation 111,14]
163. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (9) [forward demodulation 112,14]
164. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (10) [forward demodulation 116,14]
166. 14 <=> 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) [avatar definition]
169. 14 | ~13 [avatar split clause 153,149,166]
170. 14 | ~13 [avatar split clause 154,149,166]
171. 14 | ~13 [avatar split clause 155,149,166]
172. 14 | ~13 [avatar split clause 156,149,166]
173. 14 | ~9 [avatar split clause 162,105,166]
174. 14 | ~9 [avatar split clause 163,105,166]
175. 14 | ~10 [avatar split clause 164,114,166]
182. $less(X2,1) | $less(0,X2) [superposition 12,60]
193. $less(X0,$sum(X1,1)) | ~$less(X0,X1) [resolution 157,9]
195. $less(X0,$sum(1,X0)) [superposition 157,3]
200. ~$less($sum(1,X2),$sum(X2,1)) [resolution 195,20]
246. ~$less($sum(X0,1),X0) [resolution 193,8]
249. ~$less($sum(1,X6),X6) [resolution 193,200]
259. ~$less($sum($sum(X2,1),1),X2) [resolution 246,193]
268. ~$less($sum(2,X2),X2) [evaluation 259]
293. ~$less($sum(2,$sum(X2,1)),X2) [resolution 268,193]
298. ~$less($sum(X0,2),X0) [superposition 268,3]
302. ~$less($sum(3,X2),X2) [evaluation 293]
350. ~$less($sum(3,$sum(X3,1)),X3) [resolution 302,193]
354. ~$less($sum(X0,3),X0) [superposition 302,3]
357. ~$less($sum(4,X3),X3) [evaluation 350]
422. $sum($sum(-6,sK2),$sum($uminus($product(6,sK1)),X12)) = $sum(0,X12) <- (1) [superposition 4,36]
432. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (1) [superposition 4,36]
440. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (1) [superposition 36,4]
458. $sum($sum(-6,sK2),$sum($uminus($product(6,sK1)),X12)) = X12 <- (1) [forward demodulation 422,60]
493. 15 <=> 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) [avatar definition]
495. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (15) [avatar component clause 493]
496. 15 | ~1 [avatar split clause 432,34,493]
497. 15 | ~1 [avatar split clause 440,34,493]
501. $product(6,sK1) = $sum($sum(-6,sK2),0) <- (1) [superposition 458,53]
504. $sum($sum(-6,sK2),0) = $uminus($uminus($product(6,sK1))) <- (1) [superposition 458,7]
510. $product(6,sK1) = $sum(-6,sK2) <- (1) [evaluation 504]
511. $product(6,sK1) = $sum(-6,sK2) <- (1) [evaluation 501]
514. 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) <- (1, 15) [backward demodulation 495,510]
517. 16 <=> $product(6,sK1) = $sum(-6,sK2) [avatar definition]
519. $product(6,sK1) = $sum(-6,sK2) <- (16) [avatar component clause 517]
520. 16 | ~1 [avatar split clause 510,34,517]
521. 16 | ~1 [avatar split clause 511,34,517]
523. 17 <=> 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) [avatar definition]
525. 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) <- (17) [avatar component clause 523]
526. 17 | ~1 | ~15 [avatar split clause 514,493,34,523]
527. $sum(0,X0) = $sum(-6,$sum($sum(sK2,$uminus($sum(-6,sK2))),X0)) <- (17) [superposition 4,525]
528. $sum(0,X0) = $sum(-6,$sum(sK2,$sum($uminus($sum(-6,sK2)),X0))) <- (17) [forward demodulation 527,4]
529. $sum(-6,$sum(sK2,$sum($uminus($sum(-6,sK2)),X0))) = X0 <- (17) [forward demodulation 528,60]
762. $sum(-6,sK2) != $product(6,X22) | 0 = 6 | sK1 = X22 <- (16) [superposition 29,519]
773. $sum(-6,sK2) != $product(6,X22) | sK1 = X22 <- (16) [evaluation 762]
788. 0 != $sum(-6,sK2) | 0 = sK1 <- (16) [superposition 773,17]
791. 18 <=> 0 = sK1 [avatar definition]
793. 0 = sK1 <- (18) [avatar component clause 791]
795. 19 <=> 0 = $sum(-6,sK2) [avatar definition]
798. 18 | ~19 | ~16 [avatar split clause 788,517,795,791]
799. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,0),$sum(1,$uminus($product(3,sK3))))),$product(2,0)) <- (7, 18) [backward demodulation 90,793]
801. 0 = $sum($sum($uminus(sK0),$sum(1,$product(0,$product(3,0)))),$product(3,0)) <- (11, 18) [backward demodulation 133,793]
803. 0 = $sum($sum($product(9,0),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(0,$product(9,sK3)))))),$product(sK0,sK0)) <- (13, 18) [backward demodulation 151,793]
805. $sum(-6,sK2) = $product(6,0) <- (16, 18) [backward demodulation 519,793]
816. 0 = $sum(-6,sK2) <- (16, 18) [evaluation 805]
818. 0 = $sum($sum(0,$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(0,sK3))))),$product(sK0,sK0)) <- (13, 18) [evaluation 803]
820. 0 = $sum(1,$uminus(sK0)) <- (11, 18) [evaluation 801]
822. 0 = $sum(1,$sum($uminus(sK0),$sum($product(sK0,0),$uminus($product(3,sK3))))) <- (7, 18) [evaluation 799]
823. 0 = $sum(1,$sum($uminus(sK0),$sum(0,$uminus($product(3,sK3))))) <- (7, 18) [forward demodulation 822,17]
824. 0 = $sum(1,$sum($uminus(sK0),$uminus($product(3,sK3)))) <- (7, 18) [evaluation 823]
825. 0 = $sum(1,$uminus($sum($product(3,sK3),sK0))) <- (7, 18) [forward demodulation 824,6]
826. 0 = $sum(1,$uminus($sum(sK0,$product(3,sK3)))) <- (7, 18) [forward demodulation 825,3]
831. 0 = $sum($sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(0,sK3)))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 818,60]
832. 0 = $sum($sum($uminus($product(5,sK0)),$sum($uminus($product(0,sK3)),$difference(4,$product(18,sK3)))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 831,3]
833. 0 = $sum($sum($uminus($product(5,sK0)),$sum($uminus(0),$difference(4,$product(18,sK3)))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 832,123]
834. 0 = $sum($sum($uminus($product(5,sK0)),$sum(0,$difference(4,$product(18,sK3)))),$product(sK0,sK0)) <- (13, 18) [evaluation 833]
835. 0 = $sum($sum($uminus($product(5,sK0)),$difference(4,$product(18,sK3))),$product(sK0,sK0)) <- (13, 18) [forward demodulation 834,60]
846. 0 = $sum(-6,$sum(sK2,$uminus(0))) <- (16, 17, 18) [backward demodulation 525,816]
847. $sum(-6,$sum(sK2,$sum($uminus(0),X0))) = X0 <- (16, 17, 18) [backward demodulation 529,816]
849. 19 | ~16 | ~18 [avatar split clause 816,791,517,795]
850. $sum(-6,$sum(sK2,X0)) = X0 <- (16, 17, 18) [evaluation 847]
851. 0 = $sum(-6,sK2) <- (16, 17, 18) [evaluation 846]
852. 19 | ~16 | ~17 | ~18 [avatar split clause 851,791,523,517,795]
856. 20 <=> 0 = $sum(1,$uminus(sK0)) [avatar definition]
858. 0 = $sum(1,$uminus(sK0)) <- (20) [avatar component clause 856]
859. 20 | ~11 | ~18 [avatar split clause 820,791,131,856]
860. ~$less(0,$uminus(sK0)) <- (20) [superposition 249,858]
863. $less($uminus(sK0),0) <- (20) [superposition 195,858]
869. 21 <=> $less($uminus(sK0),0) [avatar definition]
871. $less($uminus(sK0),0) <- (21) [avatar component clause 869]
872. 21 | ~20 [avatar split clause 863,856,869]
873. $less($sum($uminus(sK0),X0),$sum(0,X0)) <- (21) [resolution 871,11]
875. ~$less(0,$sum($uminus(sK0),1)) <- (21) [resolution 871,20]
876. $less($sum($uminus(sK0),X0),X0) <- (21) [forward demodulation 873,60]
877. ~$less(0,$sum(1,$uminus(sK0))) <- (21) [forward demodulation 875,3]
880. 22 <=> 0 = $sum(1,$uminus($sum(sK0,$product(3,sK3)))) [avatar definition]
882. 0 = $sum(1,$uminus($sum(sK0,$product(3,sK3)))) <- (22) [avatar component clause 880]
883. 22 | ~7 | ~18 [avatar split clause 826,791,88,880]
885. 23 <=> $less(0,$uminus(sK0)) [avatar definition]
888. ~23 | ~20 [avatar split clause 860,856,885]
895. $less(0,sK0) <- (21) [superposition 876,53]
896. $less($uminus($sum(X0,sK0)),$uminus(X0)) <- (21) [superposition 876,6]
898. $less(0,$uminus($uminus(sK0))) <- (21) [superposition 876,7]
899. $less($sum(X1,$uminus(sK0)),X1) <- (21) [superposition 876,3]
901. $less(0,sK0) <- (21) [evaluation 898]
905. 24 <=> $less(0,sK0) [avatar definition]
907. $less(0,sK0) <- (24) [avatar component clause 905]
908. 24 | ~21 [avatar split clause 895,869,905]
909. 24 | ~21 [avatar split clause 901,869,905]
911. ~$less(X1,0) | $less(X1,sK0) <- (24) [resolution 907,9]
912. ~$less(sK0,$sum(0,1)) <- (24) [resolution 907,20]
913. ~$less(sK0,1) <- (24) [evaluation 912]
916. 25 <=> $less(sK0,1) [avatar definition]
918. ~$less(sK0,1) <- (~25) [avatar component clause 916]
919. ~25 | ~24 [avatar split clause 913,905,916]
921. $less(1,sK0) | 1 = sK0 <- (~25) [resolution 918,10]
934. $less($uminus($sum(sK0,$product(3,sK3))),0) <- (22) [superposition 195,882]
965. $less(X0,sK0) | $less(-1,X0) <- (24) [resolution 911,161]
969. $less($sum($uminus(sK0),0),sK0) <- (21, 24) [resolution 911,876]
970. $less($sum(0,$uminus(sK0)),sK0) <- (21, 24) [resolution 911,899]
971. $less($uminus(sK0),sK0) <- (21, 24) [resolution 911,871]
972. $less(-1,sK0) | $less(0,0) <- (24) [resolution 911,161]
973. $less(-1,sK0) <- (24) [evaluation 972]
974. $less($uminus(sK0),sK0) <- (21, 24) [forward demodulation 969,5]
975. $less($uminus(sK0),sK0) <- (21, 24) [forward demodulation 970,60]
977. 26 <=> $less(-1,sK0) [avatar definition]
979. $less(-1,sK0) <- (26) [avatar component clause 977]
980. 26 | ~24 [avatar split clause 973,905,977]
982. ~$less(X1,-1) | $less(X1,sK0) <- (26) [resolution 979,9]
983. ~$less(sK0,$sum(-1,1)) <- (26) [resolution 979,20]
984. ~$less(sK0,0) <- (26) [evaluation 983]
986. 27 <=> $less(sK0,0) [avatar definition]
989. ~27 | ~26 [avatar split clause 984,977,986]
994. 28 <=> $less($uminus(sK0),sK0) [avatar definition]
996. $less($uminus(sK0),sK0) <- (28) [avatar component clause 994]
997. 28 | ~21 | ~24 [avatar split clause 971,905,869,994]
999. 29 <=> 0 = $sum($sum($uminus($product(5,sK0)),$difference(4,$product(18,sK3))),$product(sK0,sK0)) [avatar definition]
1001. 0 = $sum($sum($uminus($product(5,sK0)),$difference(4,$product(18,sK3))),$product(sK0,sK0)) <- (29) [avatar component clause 999]
1002. 29 | ~13 | ~18 [avatar split clause 835,791,149,999]
1003. 28 | ~21 | ~24 [avatar split clause 974,905,869,994]
1004. 28 | ~21 | ~24 [avatar split clause 975,905,869,994]
1007. ~$less(sK0,$sum($uminus(sK0),1)) <- (28) [resolution 996,20]
1008. ~$less(sK0,$sum(1,$uminus(sK0))) <- (28) [forward demodulation 1007,3]
1011. 30 <=> 1 = sK0 [avatar definition]
1013. 1 = sK0 <- (30) [avatar component clause 1011]
1015. 31 <=> $less(1,sK0) [avatar definition]
1017. $less(1,sK0) <- (31) [avatar component clause 1015]
1018. 30 | 31 | 25 [avatar split clause 921,916,1015,1011]
1022. 0 = $sum(1,$uminus($sum(1,$product(3,sK3)))) <- (22, 30) [backward demodulation 882,1013]
1026. ~$less(X1,0) | $less(X1,1) <- (24, 30) [backward demodulation 911,1013]
1056. 32 <=> 0 = $sum(1,$uminus($sum(1,$product(3,sK3)))) [avatar definition]
1058. 0 = $sum(1,$uminus($sum(1,$product(3,sK3)))) <- (32) [avatar component clause 1056]
1059. 32 | ~22 | ~30 [avatar split clause 1022,1011,880,1056]
1060. $less(X0,1) | $less(-1,X0) <- (24, 30) [resolution 1026,161]
1109. $less($uminus($sum(1,$product(3,sK3))),0) <- (32) [superposition 195,1058]
1114. $less($uminus($sum(X0,1)),$uminus(X0)) <- (21, 30) [forward demodulation 896,1013]
1118. $less($uminus($sum(1,X0)),$uminus(X0)) <- (21, 30) [superposition 1114,3]
1123. $less($uminus($sum($uminus(X0),1)),X0) <- (21, 30) [superposition 1114,13]
1125. $less($uminus($sum(1,$uminus(X0))),X0) <- (21, 30) [forward demodulation 1123,3]
1130. $less($uminus(0),$uminus($uminus($sum(1,$product(3,sK3))))) <- (21, 30, 32) [superposition 1118,1058]
1136. $less(0,$sum(1,$product(3,sK3))) <- (21, 30, 32) [evaluation 1130]
1145. $less($uminus(0),$sum(1,$product(3,sK3))) <- (21, 30, 32) [superposition 1125,1058]
1147. $less(0,$sum(1,$product(3,sK3))) <- (21, 30, 32) [evaluation 1145]
1151. 33 <=> $less(0,$sum(1,$product(3,sK3))) [avatar definition]
1153. $less(0,$sum(1,$product(3,sK3))) <- (33) [avatar component clause 1151]
1154. 33 | ~21 | ~30 | ~32 [avatar split clause 1136,1056,1011,869,1151]
1155. 0 = $sum($sum($uminus($product(5,1)),$difference(4,$product(18,sK3))),$product(1,1)) <- (29, 30) [forward demodulation 1001,1013]
1156. 0 = $sum(-4,$difference(4,$product(18,sK3))) <- (29, 30) [evaluation 1155]
1157. 33 | ~21 | ~30 | ~32 [avatar split clause 1147,1056,1011,869,1151]
1160. ~$less($sum(1,$product(3,sK3)),$sum(0,1)) <- (33) [resolution 1153,20]
1161. ~$less($sum(1,$product(3,sK3)),1) <- (33) [evaluation 1160]
1170. 34 <=> $less($sum(1,$product(3,sK3)),1) [avatar definition]
1172. ~$less($sum(1,$product(3,sK3)),1) <- (~34) [avatar component clause 1170]
1173. ~34 | ~33 [avatar split clause 1161,1151,1170]
1189. $less(-1,$sum(1,$product(3,sK3))) <- (24, 30, ~34) [resolution 1172,1060]
1211. 35 <=> 0 = $sum(-4,$difference(4,$product(18,sK3))) [avatar definition]
1214. 35 | ~29 | ~30 [avatar split clause 1156,1011,999,1211]
1231. 36 <=> $less(-1,$sum(1,$product(3,sK3))) [avatar definition]
1233. $less(-1,$sum(1,$product(3,sK3))) <- (36) [avatar component clause 1231]
1234. 36 | ~24 | ~30 | 34 [avatar split clause 1189,1170,1011,905,1231]
1237. ~$less($sum(1,$product(3,sK3)),$sum(-1,1)) <- (36) [resolution 1233,20]
1238. ~$less($sum(1,$product(3,sK3)),0) <- (36) [evaluation 1237]
1259. 37 <=> $less($sum(1,$product(3,sK3)),0) [avatar definition]
1262. ~37 | ~36 [avatar split clause 1238,1231,1259]
1269. $uminus(sK2) = $sum(-6,0) <- (16, 17, 18) [superposition 850,7]
1273. -6 = $uminus(sK2) <- (16, 17, 18) [evaluation 1269]
1276. 38 <=> -6 = $uminus(sK2) [avatar definition]
1278. -6 = $uminus(sK2) <- (38) [avatar component clause 1276]
1279. 38 | ~16 | ~17 | ~18 [avatar split clause 1273,791,523,517,1276]
1283. sK2 = $uminus(-6) <- (38) [superposition 13,1278]
1289. 6 = sK2 <- (38) [evaluation 1283]
1296. 39 <=> 6 = sK2 [avatar definition]
1298. 6 = sK2 <- (39) [avatar component clause 1296]
1299. 39 | ~38 [avatar split clause 1289,1276,1296]
1306. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1))))) [forward demodulation 30,4]
1307. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)))))) [forward demodulation 1306,4]
1308. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))))))) [forward demodulation 1307,14]
1309. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)))))))) [forward demodulation 1308,14]
1310. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6))))))))) [forward demodulation 1309,14]
1311. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)))))))))) [forward demodulation 1310,14]
1312. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK0,sK2),$product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2))))))))))) [forward demodulation 1311,14]
1313. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK3,sK0)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK3,sK0),$sum(sK3,sK0)))))))))))) [forward demodulation 1312,14]
1314. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1313,3]
1315. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(1,sK3)))),$product($sum(1,sK2),$sum(sK1,1))),$uminus($sum(1,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(1,sK3)),$sum(1,sK3)),$sum(1,sK2)),$sum(1,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(1,sK3),$sum(1,sK2)),$sum(1,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum($uminus($product($product($product($product($product($product($sum(1,sK3),$sum(1,sK3)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(1,sK2),$product($sum(1,sK2),$product($sum(1,sK3),$sum(1,sK3)))))))))))) <- (30) [forward demodulation 1314,1013]
1316. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(1,sK3)))),$product($sum(1,6),$sum(sK1,1))),$uminus($sum(1,6))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(1,sK3)),$sum(1,sK3)),$sum(1,6)),$sum(1,6)),$sum(6,6)),$sum(6,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(1,sK3),$sum(1,6)),$sum(1,6)),$sum(6,6)),$sum(6,6)),$sum(6,6))),$sum($uminus($product($product($product($product($product($product($sum(1,sK3),$sum(1,sK3)),$sum(6,6)),$sum(6,6)),$sum(6,6)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(6,6),$product($sum(6,6),$product($sum(6,6),$product($sum(1,6),$product($sum(1,6),$product($sum(1,sK3),$sum(1,sK3)))))))))))) <- (30, 39) [forward demodulation 1315,1298]
1317. 0 != $product($sum($sum(-6,$sum($uminus($product(3,$sum(1,sK3))),$product(7,$sum(sK1,1)))),$product(2,$sum(sK1,1))),$sum($product($product($product($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product(84672,$sum(1,sK3))),$sum($uminus($product($product($product(1728,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product(84672,$product($sum(1,sK3),$sum(1,sK3)))))))) <- (30, 39) [evaluation 1316]
1318. 0 != $product($sum($sum(-6,$sum($uminus($product(3,$sum(1,sK3))),$product(7,$sum(sK1,1)))),$product(2,$sum(sK1,1))),$sum($product($product($product($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product($sum(sK1,1),$product(84672,$product($sum(1,sK3),$sum(1,sK3)))),$uminus($product($product($product(1728,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)),$sum(sK1,1))))))) <- (30, 39) [forward demodulation 1317,3]
1319. 0 != $product($sum($sum(-6,$sum($uminus($product(3,$sum(1,sK3))),$product(7,$sum(sK1,1)))),$product(2,$sum(sK1,1))),$sum($product($product($product($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product($sum(sK1,1),$product(84672,$product($sum(1,sK3),$sum(1,sK3)))),$uminus($product($sum(sK1,1),$product($product(1728,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)))))))) <- (30, 39) [forward demodulation 1318,14]
1320. 0 != $product($sum($sum(-6,$sum($uminus($product(3,$sum(1,sK3))),$product(7,$sum(sK1,1)))),$product(2,$sum(sK1,1))),$sum($product($product($product($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product($sum(sK1,1),$product(84672,$product($sum(1,sK3),$sum(1,sK3)))),$uminus($product($sum(sK1,1),$product($sum(sK1,1),$product(1728,$product($sum(1,sK3),$sum(1,sK3)))))))))) <- (30, 39) [forward demodulation 1319,14]
1321. 0 != $product($sum($sum(-6,$sum($uminus($product(3,$sum(1,sK3))),$product(7,$sum(1,sK1)))),$product(2,$sum(1,sK1))),$sum($product($product($product($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product($sum(1,sK1),$product(84672,$product($sum(1,sK3),$sum(1,sK3)))),$uminus($product($sum(1,sK1),$product($sum(1,sK1),$product(1728,$product($sum(1,sK3),$sum(1,sK3)))))))))) <- (30, 39) [forward demodulation 1320,3]
1322. 0 != $product($sum($sum(-6,$sum($uminus($product(3,$sum(1,sK3))),$product(7,$sum(1,0)))),$product(2,$sum(1,0))),$sum($product($product($product($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum(1,0)),$sum(1,0)),$sum(1,0)),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product($sum(1,0),$product(84672,$product($sum(1,sK3),$sum(1,sK3)))),$uminus($product($sum(1,0),$product($sum(1,0),$product(1728,$product($sum(1,sK3),$sum(1,sK3)))))))))) <- (18, 30, 39) [forward demodulation 1321,793]
1323. 0 != $product($sum(3,$uminus($product(3,$sum(1,sK3)))),$sum($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product(84672,$product($sum(1,sK3),$sum(1,sK3))),$uminus($product(1728,$product($sum(1,sK3),$sum(1,sK3)))))))) <- (18, 30, 39) [evaluation 1322]
1334. $less($uminus($sum(1,$product(3,sK3))),0) <- (22, 30) [forward demodulation 934,1013]
1336. 40 <=> $less($uminus($sum(1,$product(3,sK3))),0) [avatar definition]
1339. 40 | ~32 [avatar split clause 1109,1056,1336]
1340. 40 | ~22 | ~30 [avatar split clause 1334,1011,880,1336]
1342. 41 <=> 0 = $product($sum(3,$uminus($product(3,$sum(1,sK3)))),$sum($product(-7056,$product($sum(1,sK3),$sum(1,sK3))),$sum($uminus($product(84672,$sum(1,sK3))),$sum($product(84672,$product($sum(1,sK3),$sum(1,sK3))),$uminus($product(1728,$product($sum(1,sK3),$sum(1,sK3)))))))) [avatar definition]
1345. ~41 | ~18 | ~30 | ~39 [avatar split clause 1323,1296,1011,791,1342]
1346. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(2,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1))),$product($sum(sK1,1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1314,3]
1347. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1346,3]
1348. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($sum(1,sK1),$product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1347,14]
1349. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1348,14]
1350. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(6,sK2)),$sum(6,sK2)))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1349,14]
1351. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($sum(sK0,sK3),$sum(sK0,sK3)),$sum(6,sK2))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1350,14]
1352. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1351,14]
1353. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($sum(6,sK2),$product($product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1352,14]
1354. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1353,14]
1355. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($sum(sK0,sK3),$sum(sK0,sK2)),$sum(sK0,sK2)))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1354,14]
1356. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK2))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1355,14]
1357. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1356,14]
1358. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1357,14]
1359. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1358,14]
1360. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1359,14]
1361. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1360,14]
1362. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1361,14]
1363. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3)),$sum(sK0,sK2)))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1362,14]
1364. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK3))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1363,14]
1365. 0 != $product($sum($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(2,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product(-1,$sum(sK0,sK3)))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1364,14]
1366. 0 != $product($sum($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$product($sum(sK0,sK2),$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(2,$sum(1,sK1)))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product(-1,$sum(sK0,sK3)))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1365,4]
1367. 0 != $product($sum($sum(1,$uminus($product(3,$sum(sK0,sK3)))),$sum($product($sum(sK0,sK2),$sum(1,sK1)),$sum($uminus($sum(sK0,sK2)),$product(2,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product(-1,$sum(sK0,sK3)))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1366,4]
1368. 0 != $product($sum(1,$sum($uminus($product(3,$sum(sK0,sK3))),$sum($product($sum(sK0,sK2),$sum(1,sK1)),$sum($uminus($sum(sK0,sK2)),$product(2,$sum(1,sK1)))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product(-1,$sum(sK0,sK3)))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1367,4]
1369. 0 != $product($sum(1,$sum($uminus($product(3,$sum(sK0,sK3))),$sum($product($sum(1,sK1),$sum(sK0,sK2)),$sum($uminus($sum(sK0,sK2)),$product(2,$sum(1,sK1)))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product(-1,$sum(sK0,sK3)))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 1368,14]
1383. ~$less(X1,1) | $less(X1,sK0) <- (31) [resolution 1017,9]
1384. ~$less(sK0,$sum(1,1)) <- (31) [resolution 1017,20]
1385. ~$less(sK0,2) <- (31) [evaluation 1384]
1390. 42 <=> $less(sK0,2) [avatar definition]
1392. ~$less(sK0,2) <- (~42) [avatar component clause 1390]
1393. ~42 | ~31 [avatar split clause 1385,1015,1390]
1394. $less(2,sK0) | 2 = sK0 <- (~42) [resolution 1392,10]
1398. ~$less(0,$sum($uminus(sK0),1)) <- (21) [resolution 871,20]
1400. ~$less(0,$sum(1,$uminus(sK0))) <- (21) [forward demodulation 1398,3]
1406. ~$less(sK0,$sum($uminus(sK0),1)) <- (28) [resolution 996,20]
1407. ~$less(sK0,$sum(1,$uminus(sK0))) <- (28) [forward demodulation 1406,3]
1438. 43 <=> $less(0,$sum(1,$uminus(sK0))) [avatar definition]
1441. ~43 | ~21 [avatar split clause 877,869,1438]
1442. ~43 | ~21 [avatar split clause 1400,869,1438]
1476. $less(-1,$sum(sK0,1)) <- (24) [resolution 965,246]
1477. $less(-1,$sum(sK0,2)) <- (24) [resolution 965,298]
1478. $less(-1,$sum(sK0,3)) <- (24) [resolution 965,354]
1479. $less(-1,$sum(1,sK0)) <- (24) [resolution 965,249]
1480. $less(-1,$sum(2,sK0)) <- (24) [resolution 965,268]
1481. $less(-1,$sum(3,sK0)) <- (24) [resolution 965,302]
1482. $less(-1,$sum(4,sK0)) <- (24) [resolution 965,357]
1489. $less(-1,$sum(1,sK0)) <- (24) [forward demodulation 1476,3]
1490. $less(-1,$sum(2,sK0)) <- (24) [forward demodulation 1477,3]
1491. $less(-1,$sum(3,sK0)) <- (24) [forward demodulation 1478,3]
1493. 44 <=> $less(-1,$sum(1,sK0)) [avatar definition]
1495. $less(-1,$sum(1,sK0)) <- (44) [avatar component clause 1493]
1496. 44 | ~24 [avatar split clause 1479,905,1493]
1498. 45 <=> $less(-1,$sum(2,sK0)) [avatar definition]
1500. $less(-1,$sum(2,sK0)) <- (45) [avatar component clause 1498]
1501. 45 | ~24 [avatar split clause 1480,905,1498]
1503. 46 <=> $less(-1,$sum(3,sK0)) [avatar definition]
1505. $less(-1,$sum(3,sK0)) <- (46) [avatar component clause 1503]
1506. 46 | ~24 [avatar split clause 1481,905,1503]
1507. 44 | ~24 [avatar split clause 1489,905,1493]
1508. 45 | ~24 [avatar split clause 1490,905,1498]
1509. 46 | ~24 [avatar split clause 1491,905,1503]
1512. ~$less($sum(1,sK0),$sum(-1,1)) <- (44) [resolution 1495,20]
1513. ~$less($sum(1,sK0),0) <- (44) [evaluation 1512]
1517. ~$less($sum(2,sK0),$sum(-1,1)) <- (45) [resolution 1500,20]
1518. ~$less($sum(2,sK0),0) <- (45) [evaluation 1517]
1522. ~$less($sum(3,sK0),$sum(-1,1)) <- (46) [resolution 1505,20]
1523. ~$less($sum(3,sK0),0) <- (46) [evaluation 1522]
1526. 47 <=> $less($sum(1,sK0),0) [avatar definition]
1529. ~47 | ~44 [avatar split clause 1513,1493,1526]
1531. 48 <=> $less($sum(2,sK0),0) [avatar definition]
1534. ~48 | ~45 [avatar split clause 1518,1498,1531]
1536. 49 <=> $less($sum(3,sK0),0) [avatar definition]
1539. ~49 | ~46 [avatar split clause 1523,1503,1536]
1550. 50 <=> $less(-1,$sum(4,sK0)) [avatar definition]
1552. $less(-1,$sum(4,sK0)) <- (50) [avatar component clause 1550]
1553. 50 | ~24 [avatar split clause 1482,905,1550]
1568. ~$less($sum(4,sK0),$sum(-1,1)) <- (50) [resolution 1552,20]
1569. ~$less($sum(4,sK0),0) <- (50) [evaluation 1568]
1572. 51 <=> $less($sum(4,sK0),0) [avatar definition]
1575. ~51 | ~50 [avatar split clause 1569,1550,1572]
1582. $less($sum($uminus(sK0),-1),sK0) <- (21, 26) [resolution 982,876]
1583. $less($sum(-1,$uminus(sK0)),sK0) <- (21, 26) [resolution 982,899]
1586. $less($sum(-1,$uminus(sK0)),sK0) <- (21, 26) [forward demodulation 1582,3]
1588. 52 <=> $less(sK0,$sum(1,$uminus(sK0))) [avatar definition]
1591. ~52 | ~28 [avatar split clause 1008,994,1588]
1592. $less(X0,sK0) | $less(0,X0) <- (31) [resolution 1383,182]
1596. $less($sum($uminus(sK0),1),sK0) <- (21, 31) [resolution 1383,876]
1597. $less($sum(1,$uminus(sK0)),sK0) <- (21, 31) [resolution 1383,899]
1601. $less($sum(1,$uminus(sK0)),sK0) <- (21, 31) [forward demodulation 1596,3]
1606. ~52 | ~28 [avatar split clause 1407,994,1588]
1611. 53 <=> 2 = sK0 [avatar definition]
1615. 54 <=> $less(2,sK0) [avatar definition]
1617. $less(2,sK0) <- (54) [avatar component clause 1615]
1618. 53 | 54 | 42 [avatar split clause 1394,1390,1615,1611]
1621. ~$less(sK0,$sum(2,1)) <- (54) [resolution 1617,20]
1622. ~$less(sK0,3) <- (54) [evaluation 1621]
1624. 55 <=> $less(sK0,3) [avatar definition]
1626. ~$less(sK0,3) <- (~55) [avatar component clause 1624]
1627. ~55 | ~54 [avatar split clause 1622,1615,1624]
1628. $less(3,sK0) | 3 = sK0 <- (~55) [resolution 1626,10]
1631. 56 <=> $less($sum(-1,$uminus(sK0)),sK0) [avatar definition]
1633. $less($sum(-1,$uminus(sK0)),sK0) <- (56) [avatar component clause 1631]
1634. 56 | ~21 | ~26 [avatar split clause 1583,977,869,1631]
1635. 56 | ~21 | ~26 [avatar split clause 1586,977,869,1631]
1654. $less(0,$sum(sK0,1)) <- (31) [resolution 1592,246]
1655. $less(0,$sum(sK0,2)) <- (31) [resolution 1592,298]
1656. $less(0,$sum(sK0,3)) <- (31) [resolution 1592,354]
1657. $less(0,$sum(1,sK0)) <- (31) [resolution 1592,249]
1658. $less(0,$sum(2,sK0)) <- (31) [resolution 1592,268]
1659. $less(0,$sum(3,sK0)) <- (31) [resolution 1592,302]
1660. $less(0,$sum(4,sK0)) <- (31) [resolution 1592,357]
1671. $less(0,$sum(1,sK0)) <- (31) [forward demodulation 1654,3]
1672. $less(0,$sum(2,sK0)) <- (31) [forward demodulation 1655,3]
1673. $less(0,$sum(3,sK0)) <- (31) [forward demodulation 1656,3]
1676. 57 <=> $less(0,$sum(1,sK0)) [avatar definition]
1678. $less(0,$sum(1,sK0)) <- (57) [avatar component clause 1676]
1679. 57 | ~31 [avatar split clause 1657,1015,1676]
1681. 58 <=> $less(0,$sum(2,sK0)) [avatar definition]
1683. $less(0,$sum(2,sK0)) <- (58) [avatar component clause 1681]
1684. 58 | ~31 [avatar split clause 1658,1015,1681]
1686. 59 <=> $less(0,$sum(3,sK0)) [avatar definition]
1688. $less(0,$sum(3,sK0)) <- (59) [avatar component clause 1686]
1689. 59 | ~31 [avatar split clause 1659,1015,1686]
1690. 57 | ~31 [avatar split clause 1671,1015,1676]
1691. 58 | ~31 [avatar split clause 1672,1015,1681]
1692. 59 | ~31 [avatar split clause 1673,1015,1686]
1695. ~$less($sum(1,sK0),$sum(0,1)) <- (57) [resolution 1678,20]
1696. ~$less($sum(1,sK0),1) <- (57) [evaluation 1695]
1701. ~$less($sum(2,sK0),$sum(0,1)) <- (58) [resolution 1683,20]
1702. ~$less($sum(2,sK0),1) <- (58) [evaluation 1701]
1707. ~$less($sum(3,sK0),$sum(0,1)) <- (59) [resolution 1688,20]
1708. ~$less($sum(3,sK0),1) <- (59) [evaluation 1707]
1726. 60 <=> $less($sum(1,sK0),1) [avatar definition]
1729. ~60 | ~57 [avatar split clause 1696,1676,1726]
1731. 61 <=> $less($sum(2,sK0),1) [avatar definition]
1734. ~61 | ~58 [avatar split clause 1702,1681,1731]
1736. 62 <=> $less($sum(3,sK0),1) [avatar definition]
1739. ~62 | ~59 [avatar split clause 1708,1686,1736]
1750. 63 <=> $less(0,$sum(4,sK0)) [avatar definition]
1752. $less(0,$sum(4,sK0)) <- (63) [avatar component clause 1750]
1753. 63 | ~31 [avatar split clause 1660,1015,1750]
1756. ~$less($sum(4,sK0),$sum(0,1)) <- (63) [resolution 1752,20]
1757. ~$less($sum(4,sK0),1) <- (63) [evaluation 1756]
1761. 64 <=> $less($sum(4,sK0),1) [avatar definition]
1764. ~64 | ~63 [avatar split clause 1757,1750,1761]
1769. 65 <=> $less($sum(1,$uminus(sK0)),sK0) [avatar definition]
1771. $less($sum(1,$uminus(sK0)),sK0) <- (65) [avatar component clause 1769]
1772. 65 | ~21 | ~31 [avatar split clause 1597,1015,869,1769]
1773. 65 | ~21 | ~31 [avatar split clause 1601,1015,869,1769]
1776. ~$less(sK0,$sum($sum(-1,$uminus(sK0)),1)) <- (56) [resolution 1633,20]
1777. ~$less(sK0,$uminus(sK0)) <- (56) [evaluation 1776]
1780. 66 <=> $less(sK0,$uminus(sK0)) [avatar definition]
1783. ~66 | ~56 [avatar split clause 1777,1631,1780]
1788. ~$less(sK0,$sum($sum(1,$uminus(sK0)),1)) <- (65) [resolution 1771,20]
1789. ~$less(sK0,$sum(2,$uminus(sK0))) <- (65) [evaluation 1788]
1802. 67 <=> $less(sK0,$sum(2,$uminus(sK0))) [avatar definition]
1805. ~67 | ~65 [avatar split clause 1789,1769,1802]
1809. 68 <=> 3 = sK0 [avatar definition]
1813. 69 <=> $less(3,sK0) [avatar definition]
1816. 68 | 69 | 55 [avatar split clause 1628,1624,1813,1809]
1818. 70 <=> 0 = $product($sum(1,$sum($uminus($product(3,$sum(sK0,sK3))),$sum($product($sum(1,sK1),$sum(sK0,sK2)),$sum($uminus($sum(sK0,sK2)),$product(2,$sum(1,sK1)))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$product(-1,$sum(sK0,sK3)))))))))),$sum($uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK3))))))),$sum($uminus($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [avatar definition]
1821. ~70 [avatar split clause 1369,1818]
1822. $false [avatar sat refutation 37,42,47,52,74,83,84,85,86,91,100,101,102,103,108,117,118,134,143,144,145,146,147,152,169,170,171,172,173,174,175,496,497,520,521,526,798,849,852,859,872,883,888,908,909,919,980,989,997,1002,1003,1004,1018,1059,1154,1157,1173,1214,1234,1262,1279,1299,1339,1340,1345,1393,1441,1442,1496,1501,1506,1507,1508,1509,1529,1534,1539,1553,1575,1591,1606,1618,1627,1634,1635,1679,1684,1689,1690,1691,1692,1729,1734,1739,1753,1764,1772,1773,1783,1805,1816,1821]
% SZS output end Proof for cubes_denorm_inv3_len4_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Refutation

% Memory used [KB]: 8187
% Time elapsed: 0.368 s
% ------------------------------
% ------------------------------
% Success in time 1.875 s
