% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on cubes_denorm_inv4_len4_3
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_6 on cubes_denorm_inv4_len4_3
% Time limit reached!
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Time limit
% Termination phase: Saturation

% Memory used [KB]: 8059
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity:gtg=exists_all_2 on cubes_denorm_inv4_len4_3
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for cubes_denorm_inv4_len4_3
% SZS output start Proof for cubes_denorm_inv4_len4_3
1. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $difference($sum($product($uminus(6),1),X2),$product(6,X3)) & 0 = $sum($sum($difference($difference($difference($product(4,1),$product(18,X0)),$product($product(9,X0),X3)),$product(5,X1)),$product(9,X3)),$product(X1,X1)) & 0 = $sum($difference($sum($difference(1,$product(3,X0)),$product(X1,X3)),X1),$product(2,X3)) & 0 = $sum($difference($sum(1,$product($product(3,X3),X3)),X1),$product(3,X3)) & ~0 = $product($sum($difference($sum(1,$product($product(3,$sum(X3,1)),$sum(X3,1))),$sum(X1,X2)),$product(3,$sum(X3,1))),$sum($sum($difference($sum($product($product($product($product($product($product($product($product($uminus(1),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6))),X2),6))) [input]
2. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X3))) & 0 = $sum($sum($sum($difference($difference($product(4,1),$product(18,X0)),$product($product(9,X0),X3)),$uminus($product(5,X1))),$product(9,X3)),$product(X1,X1)) & 0 = $sum($sum($sum($difference(1,$product(3,X0)),$product(X1,X3)),$uminus(X1)),$product(2,X3)) & 0 = $sum($sum($sum(1,$product($product(3,X3),X3)),$uminus(X1)),$product(3,X3)) & ~0 = $product($sum($sum($sum(1,$product($product(3,$sum(X3,1)),$sum(X3,1))),$uminus($sum(X1,X2))),$product(3,$sum(X3,1))),$sum($sum($sum($sum($product($product($product($product($product($product($product($product($uminus(1),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$uminus($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),X2),6))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 141]
4. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 142]
5. $sum(X0,0) = X0 [theory axiom 143]
7. 0 = $sum(X0,$uminus(X0)) [theory axiom 146]
13. $uminus($uminus(X0)) = X0 [theory axiom 154]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 141]
21. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X3))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X0)),$uminus($product($product(9,X0),X3))),$uminus($product(5,X1))),$product(9,X3)),$product(X1,X1)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X0))),$product(X1,X3)),$uminus(X1)),$product(2,X3)) & 0 = $sum($sum($sum(1,$product($product(3,X3),X3)),$uminus(X1)),$product(3,X3)) & ~0 = $product($sum($sum($sum(1,$product($product(3,$sum(X3,1)),$sum(X3,1))),$uminus($sum(X1,X2))),$product(3,$sum(X3,1))),$sum($sum($sum($sum($product($product($product($product($product($product($product($product($uminus(1),$sum(X0,X1)),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1)),$product($product($product($product($product($product($sum(X1,X2),$sum(X1,X2)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X3,1)),$sum(X3,1))),$uminus($product($product($product($product($product($sum(X0,X1),$sum(X0,X1)),$sum(X1,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),X2),6))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X1))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X3)),$uminus($product($product(9,X3),X1))),$uminus($product(5,X0))),$product(9,X1)),$product(X0,X0)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X3))),$product(X0,X1)),$uminus(X0)),$product(2,X1)) & 0 = $sum($sum($sum(1,$product($product(3,X1),X1)),$uminus(X0)),$product(3,X1)) & ~0 = $product($sum($sum($sum(1,$product($product(3,$sum(X1,1)),$sum(X1,1))),$uminus($sum(X0,X2))),$product(3,$sum(X1,1))),$sum($sum($sum($sum($product($product($product($product($product($product($product($product($uminus(1),$sum(X3,X0)),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$product($product($product($product($product($product($sum(X0,X2),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1))),$uminus($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),X2),6))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum($product($uminus(6),1),X2),$uminus($product(6,X1))) & 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,X3)),$uminus($product($product(9,X3),X1))),$uminus($product(5,X0))),$product(9,X1)),$product(X0,X0)) & 0 = $sum($sum($sum($sum(1,$uminus($product(3,X3))),$product(X0,X1)),$uminus(X0)),$product(2,X1)) & 0 = $sum($sum($sum(1,$product($product(3,X1),X1)),$uminus(X0)),$product(3,X1)) & 0 != $product($sum($sum($sum(1,$product($product(3,$sum(X1,1)),$sum(X1,1))),$uminus($sum(X0,X2))),$product(3,$sum(X1,1))),$sum($sum($sum($sum($product($product($product($product($product($product($product($product($uminus(1),$sum(X3,X0)),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1)),$product($product($product($product($product($product($sum(X0,X2),$sum(X0,X2)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X1,1)),$sum(X1,1))),$uminus($product($product($product($product($product($sum(X3,X0),$sum(X3,X0)),$sum(X0,X2)),$sum(X2,6)),$sum(X2,6)),$sum(X2,6)))),X2),6))) [flattening 22]
24. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($sum($sum($product($product($product($product($product($product($product($product($uminus(1),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))),sK2),6)) [cnf transformation 23]
25. 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) [cnf transformation 23]
26. 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) [cnf transformation 23]
27. 0 = $sum($sum($sum($sum($difference($product(4,1),$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [cnf transformation 23]
28. 0 = $sum($sum($product($uminus(6),1),sK2),$uminus($product(6,sK1))) [cnf transformation 23]
30. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($sum($sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))),sK2),6)) [evaluation 24]
31. 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [evaluation 27]
32. 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) [evaluation 28]
34. 1 <=> 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) [avatar definition]
36. 0 = $sum($sum(-6,sK2),$uminus($product(6,sK1))) <- (1) [avatar component clause 34]
37. 1 [avatar split clause 32,34]
39. 2 <=> 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) [avatar definition]
41. 0 = $sum($sum($sum(1,$product($product(3,sK1),sK1)),$uminus(sK0)),$product(3,sK1)) <- (2) [avatar component clause 39]
42. 2 [avatar split clause 25,39]
44. 3 <=> 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) [avatar definition]
46. 0 = $sum($sum($sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1)),$uminus(sK0)),$product(2,sK1)) <- (3) [avatar component clause 44]
47. 3 [avatar split clause 26,44]
49. 4 <=> 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) [avatar definition]
51. 0 = $sum($sum($sum($sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))),$uminus($product(5,sK0))),$product(9,sK1)),$product(sK0,sK0)) <- (4) [avatar component clause 49]
52. 4 [avatar split clause 31,49]
53. 0 = $sum($uminus(X0),X0) [superposition 7,13]
57. 0 = $sum($sum($uminus(sK0),$sum($sum(1,$uminus($product(3,sK3))),$product(sK0,sK1))),$product(2,sK1)) <- (3) [backward demodulation 46,3]
58. 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) <- (2) [backward demodulation 41,3]
59. 0 = $sum($sum($sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))),$product(9,sK1)),$product(sK0,sK0)) <- (4) [backward demodulation 51,3]
60. $sum(0,X0) = X0 [superposition 3,5]
68. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) <- (3) [forward demodulation 57,3]
69. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) <- (4) [forward demodulation 59,3]
71. 5 <=> 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) [avatar definition]
73. 0 = $sum($sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1))),$product(3,sK1)) <- (5) [avatar component clause 71]
74. 5 | ~2 [avatar split clause 58,39,71]
75. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 73,3]
76. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 73,3]
77. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 3,73]
78. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (5) [superposition 3,73]
80. 6 <=> 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) [avatar definition]
82. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product($product(3,sK1),sK1)))) <- (6) [avatar component clause 80]
83. 6 | ~5 [avatar split clause 75,71,80]
84. 6 | ~5 [avatar split clause 76,71,80]
85. 6 | ~5 [avatar split clause 77,71,80]
86. 6 | ~5 [avatar split clause 78,71,80]
88. 7 <=> 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) [avatar definition]
90. 0 = $sum($sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3))))),$product(2,sK1)) <- (7) [avatar component clause 88]
91. 7 | ~3 [avatar split clause 68,44,88]
92. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 90,3]
93. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 90,3]
94. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 3,90]
95. 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) <- (7) [superposition 3,90]
97. 8 <=> 0 = $sum($product(2,sK1),$sum($uminus(sK0),$sum($product(sK0,sK1),$sum(1,$uminus($product(3,sK3)))))) [avatar definition]
100. 8 | ~7 [avatar split clause 92,88,97]
101. 8 | ~7 [avatar split clause 93,88,97]
102. 8 | ~7 [avatar split clause 94,88,97]
103. 8 | ~7 [avatar split clause 95,88,97]
105. 9 <=> 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) [avatar definition]
107. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1))))),$product(sK0,sK0)) <- (9) [avatar component clause 105]
108. 9 | ~4 [avatar split clause 69,49,105]
109. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 107,3]
110. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 107,3]
111. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 3,107]
112. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (9) [superposition 3,107]
114. 10 <=> 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) [avatar definition]
116. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product($product(9,sK3),sK1)))))) <- (10) [avatar component clause 114]
117. 10 | ~9 [avatar split clause 109,105,114]
118. 10 | ~9 [avatar split clause 110,105,114]
119. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (6) [backward demodulation 82,14]
120. 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) <- (5) [backward demodulation 73,14]
121. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) <- (9) [backward demodulation 107,14]
131. 11 <=> 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) [avatar definition]
133. 0 = $sum($sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1)))),$product(3,sK1)) <- (11) [avatar component clause 131]
134. 11 | ~5 [avatar split clause 120,71,131]
135. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 133,3]
136. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 133,3]
137. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 3,133]
138. 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) <- (11) [superposition 3,133]
140. 12 <=> 0 = $sum($product(3,sK1),$sum($uminus(sK0),$sum(1,$product(sK1,$product(3,sK1))))) [avatar definition]
143. 12 | ~6 [avatar split clause 119,80,140]
144. 12 | ~11 [avatar split clause 135,131,140]
145. 12 | ~11 [avatar split clause 136,131,140]
146. 12 | ~11 [avatar split clause 137,131,140]
147. 12 | ~11 [avatar split clause 138,131,140]
149. 13 <=> 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) [avatar definition]
151. 0 = $sum($sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3)))))),$product(sK0,sK0)) <- (13) [avatar component clause 149]
152. 13 | ~9 [avatar split clause 121,105,149]
153. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 151,3]
154. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 151,3]
155. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 3,151]
156. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (13) [superposition 3,151]
162. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (9) [forward demodulation 111,14]
163. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (9) [forward demodulation 112,14]
164. 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) <- (10) [forward demodulation 116,14]
166. 14 <=> 0 = $sum($product(sK0,sK0),$sum($product(9,sK1),$sum($uminus($product(5,sK0)),$sum($difference(4,$product(18,sK3)),$uminus($product(sK1,$product(9,sK3))))))) [avatar definition]
169. 14 | ~13 [avatar split clause 153,149,166]
170. 14 | ~13 [avatar split clause 154,149,166]
171. 14 | ~13 [avatar split clause 155,149,166]
172. 14 | ~13 [avatar split clause 156,149,166]
173. 14 | ~9 [avatar split clause 162,105,166]
174. 14 | ~9 [avatar split clause 163,105,166]
175. 14 | ~10 [avatar split clause 164,114,166]
422. $sum($sum(-6,sK2),$sum($uminus($product(6,sK1)),X12)) = $sum(0,X12) <- (1) [superposition 4,36]
432. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (1) [superposition 4,36]
440. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (1) [superposition 36,4]
458. $sum($sum(-6,sK2),$sum($uminus($product(6,sK1)),X12)) = X12 <- (1) [forward demodulation 422,60]
493. 15 <=> 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) [avatar definition]
495. 0 = $sum(-6,$sum(sK2,$uminus($product(6,sK1)))) <- (15) [avatar component clause 493]
496. 15 | ~1 [avatar split clause 432,34,493]
497. 15 | ~1 [avatar split clause 440,34,493]
501. $product(6,sK1) = $sum($sum(-6,sK2),0) <- (1) [superposition 458,53]
504. $sum($sum(-6,sK2),0) = $uminus($uminus($product(6,sK1))) <- (1) [superposition 458,7]
510. $product(6,sK1) = $sum(-6,sK2) <- (1) [evaluation 504]
511. $product(6,sK1) = $sum(-6,sK2) <- (1) [evaluation 501]
514. 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) <- (1, 15) [backward demodulation 495,510]
517. 16 <=> $product(6,sK1) = $sum(-6,sK2) [avatar definition]
520. 16 | ~1 [avatar split clause 510,34,517]
521. 16 | ~1 [avatar split clause 511,34,517]
523. 17 <=> 0 = $sum(-6,$sum(sK2,$uminus($sum(-6,sK2)))) [avatar definition]
526. 17 | ~1 | ~15 [avatar split clause 514,493,34,523]
655. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))),$sum(sK2,6))) [forward demodulation 30,4]
656. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1))),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum(sK2,6)))) [forward demodulation 655,4]
657. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))),$sum(sK2,6))))) [forward demodulation 656,4]
658. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($sum(sK2,6),$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6))))))) [forward demodulation 657,3]
659. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)))))))) [forward demodulation 658,4]
660. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($sum(sK2,6),$product($product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6))))))))) [forward demodulation 659,14]
661. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($sum(sK2,6),$product($sum(sK2,6),$product($product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK2,6)))))))))) [forward demodulation 660,14]
662. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($product($sum(sK3,sK0),$sum(sK3,sK0)),$sum(sK0,sK2))))))))))) [forward demodulation 661,14]
663. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK3,sK0)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK0,sK2),$product($sum(sK3,sK0),$sum(sK3,sK0)))))))))))) [forward demodulation 662,14]
664. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(sK2,6)),$sum(sK2,6)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK2,6),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 663,3]
665. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 664,3]
666. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($sum(sK1,1),$product($product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 665,14]
667. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($sum(sK1,1),$product($sum(sK1,1),$product($product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 666,14]
668. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(6,sK2),$product($product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 667,14]
669. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($sum(sK0,sK2),$sum(sK0,sK2)),$sum(sK0,sK2)))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 668,14]
670. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(sK1,1)),$sum(sK1,1))),$uminus($sum(sK0,sK2))),$product(3,$sum(sK1,1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(sK1,1)),$sum(sK1,1)),$sum($product($sum(sK1,1),$product($sum(sK1,1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 669,14]
671. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1)),$sum(1,sK1)),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 670,3]
672. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 671,14]
673. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2)),$sum(6,sK2)))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 672,14]
674. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)),$sum(6,sK2))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 673,14]
675. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2)),$sum(6,sK2)))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 674,14]
676. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)),$sum(sK0,sK2))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 675,14]
677. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($product(-1,$sum(sK0,sK3)),$sum(sK0,sK2)))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 676,14]
678. 0 != $product($sum($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$uminus($sum(sK0,sK2))),$product(3,$sum(1,sK1))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product(-1,$sum(sK0,sK3))))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 677,14]
679. 0 != $product($sum($sum(1,$product($product(3,$sum(1,sK1)),$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1)))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product(-1,$sum(sK0,sK3))))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 678,4]
680. 0 != $product($sum(1,$sum($product($product(3,$sum(1,sK1)),$sum(1,sK1)),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product(-1,$sum(sK0,sK3))))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 679,4]
681. 0 != $product($sum(1,$sum($product($sum(1,sK1),$product(3,$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product(-1,$sum(sK0,sK3))))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [forward demodulation 680,14]
709. 18 <=> 0 = $product($sum(1,$sum($product($sum(1,sK1),$product(3,$sum(1,sK1))),$sum($uminus($sum(sK0,sK2)),$product(3,$sum(1,sK1))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$product(-1,$sum(sK0,sK3))))))))),$sum($product($sum(1,sK1),$product($sum(1,sK1),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK2),$sum(sK0,sK2))))))),$sum(sK2,$sum(6,$uminus($product($sum(6,sK2),$product($sum(6,sK2),$product($sum(6,sK2),$product($sum(sK0,sK2),$product($sum(sK0,sK3),$sum(sK0,sK3)))))))))))) [avatar definition]
712. ~18 [avatar split clause 681,709]
713. $false [avatar sat refutation 37,42,47,52,74,83,84,85,86,91,100,101,102,103,108,117,118,134,143,144,145,146,147,152,169,170,171,172,173,174,175,496,497,520,521,526,712]
% SZS output end Proof for cubes_denorm_inv4_len4_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 7537e19c0 on 2021-06-08 12:31:04 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Refutation

% Memory used [KB]: 7547
% Time elapsed: 0.080 s
% ------------------------------
% ------------------------------
% Success in time 1.586 s
