% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on eucliddiv_norm_inv2_1xlen2
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for eucliddiv_norm_inv2_1xlen2
% SZS output start Proof for eucliddiv_norm_inv2_1xlen2
1. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($difference($difference(X1,r00),$product(q00,y00)),$product(X0,y00)) & ~0 = $sum($sum($sum($product($product($product($product(0,$difference(X1,y00)),$difference(X1,y00)),$difference(X1,y00)),y00),$product($product($product($product($product(0,$sum(X0,1)),$difference(X1,y00)),$difference(X1,y00)),y00),y00)),$product($product($product($product($product(0,q00),$difference(X1,y00)),$difference(X1,y00)),y00),y00)),$product($product($product($product(0,$difference(X1,y00)),$difference(X1,y00)),r00),y00))) [input]
2. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($sum($difference(X1,r00),$uminus($product(q00,y00))),$product(X0,y00)) & ~0 = $sum($sum($sum($product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(X0,1)),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),r00),y00))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 143]
4. $sum(X0,$sum(X1,X2)) = $sum($sum(X0,X1),X2) [theory axiom 144]
5. $sum(X0,0) = X0 [theory axiom 145]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 143]
15. $product(X0,$product(X1,X2)) = $product($product(X0,X1),X2) [theory axiom 144]
17. 0 = $product(X0,0) [theory axiom 157]
18. $product(X0,$sum(X1,X2)) = $sum($product(X0,X1),$product(X0,X2)) [theory axiom 158]
21. ? [X2 : $int,X1 : $int,X0 : $int] : (0 = $sum($product($uminus(1),y00),X2) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X0,y00)) & ~0 = $sum($sum($sum($product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(X0,1)),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),r00),y00))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int] : (0 = $sum($product($uminus(1),y00),X0) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X2,y00)) & ~0 = $sum($sum($sum($product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(X2,1)),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),r00),y00))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int] : (0 = $sum($product($uminus(1),y00),X0) & 0 = $sum($sum($sum(X1,$uminus(r00)),$uminus($product(q00,y00))),$product(X2,y00)) & 0 != $sum($sum($sum($product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(X2,1)),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),y00),y00)),$product($product($product($product(0,$sum(X1,$uminus(y00))),$sum(X1,$uminus(y00))),r00),y00))) [flattening 22]
24. 0 != $sum($sum($sum($product($product($product($product(0,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product(0,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),r00),y00)) [cnf transformation 23]
38. $sum(0,X0) = X0 [superposition 3,5]
57. 0 = $product(0,X1) [superposition 14,17]
938. 0 != $sum($sum($sum($product($product($product($product(0,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product(y00,$product($product($product(0,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),r00))) [forward demodulation 24,14]
939. 0 != $sum($sum($sum($product($product($product($product(0,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product(y00,$product(r00,$product($product(0,$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00)))))) [forward demodulation 938,14]
940. 0 != $sum($sum($sum($product($product($product($sum(sK1,$uminus(y00)),$product(0,$sum(sK1,$uminus(y00)))),$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product(y00,$product(r00,$product($sum(sK1,$uminus(y00)),$product(0,$sum(sK1,$uminus(y00))))))) [forward demodulation 939,14]
941. 0 != $sum($sum($sum($product($product($product($sum(sK1,$uminus(y00)),0),$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product(y00,$product(r00,$product($sum(sK1,$uminus(y00)),0)))) [forward demodulation 940,57]
942. 0 != $sum($sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product(y00,$product(r00,0))) [forward demodulation 941,17]
943. 0 != $sum($sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product(y00,0)) [forward demodulation 942,17]
944. 0 != $sum($sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),0) [forward demodulation 943,17]
945. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product($product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00))),0) [forward demodulation 944,3]
946. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product($product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00)))),0) [forward demodulation 945,14]
947. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))))))),0) [forward demodulation 946,14]
948. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product($product(0,$sum(sK2,1)),$sum(sK1,$uminus(y00)))))))),0) [forward demodulation 947,14]
949. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product($sum(sK1,$uminus(y00)),$product(0,$sum(sK2,1)))))))),0) [forward demodulation 948,14]
950. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product($sum(sK1,$uminus(y00)),0)))))),0) [forward demodulation 949,57]
951. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),0))))),0) [forward demodulation 950,17]
952. 0 != $sum($sum($product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00),$sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,0)))),0) [forward demodulation 951,17]
953. 0 != $sum($sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,0))),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00)),0) [forward demodulation 952,3]
954. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$sum($product(y00,$product(y00,0)),$product($product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00),y00))),0) [forward demodulation 953,4]
955. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$sum($product(y00,$product(y00,0)),$product(y00,$product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00)))),0) [forward demodulation 954,14]
956. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$sum($product(y00,0),$product($product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))),y00)))),0) [forward demodulation 955,18]
957. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$sum($product(y00,0),$product(y00,$product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))))))),0) [forward demodulation 956,14]
958. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$sum(0,$product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00))))))),0) [forward demodulation 957,18]
959. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($product($product(0,q00),$sum(sK1,$uminus(y00))),$sum(sK1,$uminus(y00)))))),0) [forward demodulation 958,38]
960. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product($product(0,q00),$sum(sK1,$uminus(y00))))))),0) [forward demodulation 959,14]
961. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),$product(0,$product(q00,$sum(sK1,$uminus(y00)))))))),0) [forward demodulation 960,15]
962. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,$product($sum(sK1,$uminus(y00)),0)))),0) [forward demodulation 961,57]
963. 0 != $sum($sum($product($product(0,$sum(sK1,$uminus(y00))),y00),$product(y00,$product(y00,0))),0) [forward demodulation 962,17]
964. 0 != $sum($sum($product(y00,$product(y00,0)),$product($product(0,$sum(sK1,$uminus(y00))),y00)),0) [forward demodulation 963,3]
965. 0 != $sum($sum($product(y00,$product(y00,0)),$product(y00,$product(0,$sum(sK1,$uminus(y00))))),0) [forward demodulation 964,14]
966. 0 != $sum($product(y00,$sum($product(y00,0),$product(0,$sum(sK1,$uminus(y00))))),0) [forward demodulation 965,18]
967. 0 != $sum($product(y00,$sum($product(y00,0),0)),0) [forward demodulation 966,57]
968. 0 != $sum($product(y00,$product(y00,0)),0) [forward demodulation 967,5]
969. 0 != $sum($product(y00,0),0) [forward demodulation 968,17]
970. 0 != $sum(0,0) [forward demodulation 969,17]
971. $false [evaluation 970]
% SZS output end Proof for eucliddiv_norm_inv2_1xlen2
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Refutation

% Memory used [KB]: 7803
% Time elapsed: 0.059 s
% ------------------------------
% ------------------------------
% Success in time 0.064 s
