% Warning: check-sat is not the last entry. Skipping the rest!
% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on simpleloopV2_inv1
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for simpleloopV2_inv1
% SZS output start Proof for simpleloopV2_inv1
1. ? [X1 : $int,X0 : $int] : (0 = $sum($difference($difference(X0,y00),x00),X1) & 0 = $sum($difference($difference(X0,y00),x00),X1) & (~0 = $sum($sum($difference($difference($difference(X0,1),y00),x00),X1),1) | ~0 = $sum($difference($difference($sum(X0,1),y00),x00),$difference(X1,1)))) [input]
2. ? [X1 : $int,X0 : $int] : (0 = $sum($sum($difference(X0,y00),$uminus(x00)),X1) & 0 = $sum($sum($difference(X0,y00),$uminus(x00)),X1) & (~0 = $sum($sum($sum($difference($difference(X0,1),y00),$uminus(x00)),X1),1) | ~0 = $sum($sum($difference($sum(X0,1),y00),$uminus(x00)),$sum(X1,$uminus(1))))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 143]
15. ? [X1 : $int,X0 : $int] : (0 = $sum($sum($sum(X0,$uminus(y00)),$uminus(x00)),X1) & 0 = $sum($sum($sum(X0,$uminus(y00)),$uminus(x00)),X1) & (~0 = $sum($sum($sum($sum($difference(X0,1),$uminus(y00)),$uminus(x00)),X1),1) | ~0 = $sum($sum($sum($sum(X0,1),$uminus(y00)),$uminus(x00)),$sum(X1,$uminus(1))))) [theory normalization 2]
16. ? [X0 : $int,X1 : $int] : (0 = $sum($sum($sum(X1,$uminus(y00)),$uminus(x00)),X0) & 0 = $sum($sum($sum(X1,$uminus(y00)),$uminus(x00)),X0) & (~0 = $sum($sum($sum($sum($difference(X1,1),$uminus(y00)),$uminus(x00)),X0),1) | ~0 = $sum($sum($sum($sum(X1,1),$uminus(y00)),$uminus(x00)),$sum(X0,$uminus(1))))) [rectify 15]
17. ? [X0 : $int,X1 : $int] : (0 = $sum($sum($sum(X1,$uminus(y00)),$uminus(x00)),X0) & 0 = $sum($sum($sum(X1,$uminus(y00)),$uminus(x00)),X0) & (0 != $sum($sum($sum($sum($difference(X1,1),$uminus(y00)),$uminus(x00)),X0),1) | 0 != $sum($sum($sum($sum(X1,1),$uminus(y00)),$uminus(x00)),$sum(X0,$uminus(1))))) [flattening 16]
18. 0 != $sum($sum($sum($sum(sK1,1),$uminus(y00)),$uminus(x00)),$sum(sK0,$uminus(1))) | 0 != $sum($sum($sum($sum($difference(sK1,1),$uminus(y00)),$uminus(x00)),sK0),1) [cnf transformation 17]
19. 0 = $sum($sum($sum(sK1,$uminus(y00)),$uminus(x00)),sK0) [cnf transformation 17]
20. 0 = $sum($sum($sum(sK1,$uminus(y00)),$uminus(x00)),sK0) [cnf transformation 17]
21. 0 != $sum(sK1,$sum($uminus(y00),$sum($uminus(x00),sK0))) | 0 != $sum($sum($sum($sum($difference(sK1,1),$uminus(y00)),$uminus(x00)),sK0),1) [evaluation 18]
380. 0 = $sum($sum($uminus(x00),$sum(sK1,$uminus(y00))),sK0) [forward demodulation 19,3]
426. 0 = $sum($sum($uminus(x00),$sum(sK1,$uminus(y00))),sK0) [forward demodulation 20,3]
473. 1 <=> 0 = $sum($sum($uminus(x00),$sum(sK1,$uminus(y00))),sK0) [avatar definition]
476. 1 [avatar split clause 380,473]
522. 1 [avatar split clause 426,473]
728. 0 != $sum(sK1,$sum($uminus(y00),$sum(sK0,$uminus(x00)))) | 0 != $sum($sum($sum($sum($difference(sK1,1),$uminus(y00)),$uminus(x00)),sK0),1) [forward demodulation 21,3]
729. 0 != $sum($sum(sK0,$sum($sum($difference(sK1,1),$uminus(y00)),$uminus(x00))),1) | 0 != $sum(sK1,$sum($uminus(y00),$sum(sK0,$uminus(x00)))) [forward demodulation 728,3]
730. 0 != $sum($sum(sK0,$sum($uminus(x00),$sum($difference(sK1,1),$uminus(y00)))),1) | 0 != $sum(sK1,$sum($uminus(y00),$sum(sK0,$uminus(x00)))) [forward demodulation 729,3]
731. 0 != $sum($sum(sK0,$sum($uminus(x00),$sum($uminus(y00),$difference(sK1,1)))),1) | 0 != $sum(sK1,$sum($uminus(y00),$sum(sK0,$uminus(x00)))) [forward demodulation 730,3]
781. 2 <=> 0 = $sum(sK1,$sum($uminus(y00),$sum(sK0,$uminus(x00)))) [avatar definition]
785. 3 <=> 0 = $sum($sum(sK0,$sum($uminus(x00),$sum($uminus(y00),$difference(sK1,1)))),1) [avatar definition]
788. ~2 | ~3 [avatar split clause 731,785,781]
789. $false [avatar smt refutation 476,522,788]
% SZS output end Proof for simpleloopV2_inv1
% ------------------------------
% Version: Vampire 4.6.1 (commit dd79494d8 on 2022-01-05 14:29:51 +0000)
% Linked with Z3 4.8.15.0 dfe2b27f9aa6d72583014fa301bd79d223fb254f z3-4.8.4-5822-gdfe2b27f9
% Termination reason: Refutation

% Memory used [KB]: 7675
% Time elapsed: 0.057 s
% ------------------------------
% ------------------------------
% Success in time 0.059 s
